# Programming in Python for Data Science 

# Assignment 3: Tidy Data and Joining Dataframes

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are the links of 2 videos that explain:

1. [How to save your Jupyter notebook](https://youtu.be/0aoLgBoAUSA) and,       
2. [How to answer a question in a Jupyter notebook assignment](https://youtu.be/7j0WKhI3W4s).

### Assignment Learning Goals:

By the end of the module, students are expected to:

- Explain what tidy data is.
- Use [`.melt()`](https://pandas.pydata.org/docs/reference/api/pandas.melt.html) and [`.pivot()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.pivot.html) to reshape dataframes, specifically to make tidy data.
- Learn how to reset a dataframe's index.
- Combine dataframes using [`.merge()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.merge.html) and [`.concat()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.concat.html) and know when to use these different methods.
- Understand the different joining methods.

This assignment covers [Module 3](https://prog-learn.mds.ubc.ca/en/module3) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` and the `raise NotImplementedError # No Answer - remove if you provide an answer` with your completed code and answers then proceed to run the cell!

Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct. These questions are worth 2 points.

```python
# Import libraries needed for this lab
from hashlib import sha1
import pandas as pd
import test_assignment3 as t
import altair as alt
```

## 1.  Tidy Data

**Question 1(a)** <br> {points: 1}  

Which of the following do **NOT** characterize a tidy dataset? *note - there may be more than 1 correct answers to this question*

A) Each row is a single variable

B) There are no missing or erroneous values

C) Each value is a single cell

D) Each variable is a single column

_Assign your answer to an object called `answer1_a` in the code chunk below. Make sure your answer contains uppercase letters and surround it with quotation marks and square brackets.  If there are more than one answers to this question, separate each letter with a comma within the square brackets. For example if you believe the answer is A, B and C your answer would like this:      
`answer1_a = ['A', 'B', 'C')`_


<!-- grade_id: cell-885c670d16b8075a, points: 0, grade: false, solution: true -->
```python
answer1_a = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b9575c023870ca02, points: 1, grade: true, solution: false -->
```python
t.test_1a(answer1_a)
```

**Question 1(b)** 
<br> {points: 1}  

The data below is wine ratings given for 3 wines by 5 different wine tasters. We are interested in seeing if Taster or Wine type influences the rating.  Given that motivation, which arrangement of the data set show below is "tidy"?

##### Data set 1:

|     Taster       | Chardonnay | Pinot Grigio | Pinot Blanc |
|------------|------------|--------------|-----------------|
| 001 | 75         | 89           | 92              |
| 002 | 89         | 88           | 89              |
| 003 | 72         | 90           | 95              |
| 004 | 85         | 81           | 90              |
| 005 | 83         | 89           | 88              |


##### Data set 2:
|   Wine | Taster 001 | Taster 002 | Taster 003 | Taster 004 | Taster 005 |
|------------|------------|--------------|-----------------|-------|---------|
| Chardonnay | 75         | 89           | 72              | 85 | 83
| Pinot Grigio | 89         | 88           | 90             | 81 | 89 |
| Pinot Blanc | 92         | 89           | 95              | 90 | 88 |

##### Data set 3:

| Taster           | Wine | Rating | 
|------------|------------|----|
| 001 |  Chardonnay |  75         | 
| 002 |  Chardonnay | 89         | 
| 003 |  Chardonnay |72         |  
| 004 |  Chardonnay |85         | 
| 005 | Chardonnay | 83         |  
| 001 |  Pinot Grigio | 89         | 
| 002 |  Pinot Grigio | 88         | 
| 003 |  Pinot Grigio | 90         | 
| 004 |  Pinot Grigio | 81         | 
| 005 |  Pinot Grigio | 90         | 
| 001 |  Pinot Blanc | 92         | 
| 002 | Pinot Blanc | 89         |
| 003 | Pinot Blanc | 95         | 
| 004 | Pinot Blanc | 90         | 
| 005 | Pinot Blanc | 88         | 

##### Data set 4: 
| Taster    | Chardonnay Rating | 
|------------|------------|
| 001 |  75         | 
| 002 |   89         | 
| 003 |  72         |  
| 004 | 85         | 
| 005 | 83         |  

| Taster           | Pinot Grigio Rating | 
|------------|------------|
| 001 |   89         | 
| 002 |  88         | 
| 003 |  90         | 
| 004 | 81         | 
| 005 |  90         | 

| Taster           | Pinot Blanc Rating | 
|------------|------------|
| 001 |   92         | 
| 002 |  89         |
| 003 |  95         | 
| 004 |  90         | 
| 005 |  88         | 

Given that motivation, which arrangement of the data set show below is "tidy"?

_Assign your answer to an object called `answer1_b` in the code chunk below. Make sure your answer is surrounded by square brackets. If there are more than one answers to this question, separate each number with a comma in the square brackets. For example if you believe the answer is 1, 2 and 3 your answer would like this:   
`answer1_b = [1, 2, 3]`_



<!-- grade_id: cell-056483e434ae664a, points: 0, grade: false, solution: true -->
```python
answer1_b = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e35eba9f97d1d7bf, points: 1, grade: true, solution: false -->
```python
t.test_1b(answer1_b)
```

**Question 1(c)** <br> {points: 2}  

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   

Why is the primary goal of data wrangling getting dataframes into the tidy data format?
 
A) Having data expressed in such a way, allows for easier readability and is more aesthetically pleasing. 

B) Tidy format uses less storage space on your computer. 

C) Many or most modern Data Science tools accept the tidy data format directly (or very close to that) and we need to get the data in a state ready for analysis.

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer1_c`.*


<!-- grade_id: cell-c5a440fc14cdd03d, points: 0, grade: false, solution: true -->
```python
answer1_c = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-8277a8851358fbf9, points: 2, grade: true, solution: false -->
```python
assert 'answer1_c' in locals(), "Your answer should be named 'answer1_c'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
```

**Question 1(d)** <br> {points: 1}  

Use `.pivot()` or `.melt()` to tidy the `planes1` dataframe below to convert it to tidy data. Name the new dataframe `planes2`


```python
flights = {'on_time': ['yes', 'no'],
           'airline_US': [18, 17],
           'airline_UA' : [13, 9]
          }
planes1 = pd.DataFrame(flights, columns = ['on_time', 'airline_US', 'airline_UA'])
planes1

```

<!-- grade_id: cell-2d6f65755b4e202b, points: 0, grade: false, solution: true -->
```python
planes2 = None
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-ec7481e5b9a9f30a, points: 1, grade: true, solution: false -->
```python
t.test_1d(planes2)
```

**Question 1(e)** <br> {points: 1}  

We are interested in determining if the number of medals is influenced by the country and the type of medal. Given this statistical question, is the below dataset tidy data? 

\begin{array}{ccccc}
& & COUNTRY & GOLD & SILVER & BRONZE\\
& 1 & USA & 46 & 37 & 38\\
medals\_df = & 2 & CHN & 38 & 23 & 17\\
& 3 & GBR & 29 & 18 & 26\\
\end{array}


To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
 
A) The dataframe is tidy.

B) The dataframe is ***NOT*** tidy and cannot be changed to tidy data by using `.pivot()` or `.melt()`. 

C) The dataframe is ***NOT*** tidy and can be changed to tidy data by using `.pivot()`.

D) The dataframe is ***NOT*** tidy and can be changed to tidy data by using `.melt()`.

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer1_e`.*

<!-- grade_id: cell-492bb19dc3a7130f, points: 0, grade: false, solution: true -->
```python
answer1_e = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-01a6e527e9db7a0e, points: 1, grade: true, solution: false -->
```python
t.test_1e(answer1_e)
```

## 2.  Pivot and Melt

**Question 2(a)** <br> {points: 1}  

The file at this [url](https://raw.githubusercontent.com/plotly/datasets/master/school_earnings.csv) contains data about the average wage in 1000 dollars per year for men and women graduates from various US universities. 

***Note: many popular datasets have sex as a feature where the possible values are male and female. This representation reflects how the data were collected and is not meant to imply that, for example, gender is binary.***


Load only the `School`, `Women` and `Men` columns into a dataframe named `wages`.



<!-- grade_id: cell-6e9f73a26ab62a99, points: 0, grade: false, solution: true -->
```python
wages = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e0cd5c037882035a, points: 1, grade: true, solution: false -->
```python
t.test_2a(wages)
```

**Question 2(b)** <br> {points: 1}  

We are interested in seeing if a person's wage is influenced by the variables gender and school. Given this statistical question, which tidy data criterion is not being satisfied in the dataframe?

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
 
A) Criterion 1: Each row is a single observation      

B) Criterion 2: Each variable is a single column     

C) Criterion 3: Each value is a single cell      

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer2_b`.*

<!-- grade_id: cell-603c14df9149a8a2, points: 0, grade: false, solution: true -->
```python
answer2_b = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-239c2fb98499b9f4, points: 1, grade: true, solution: false -->
```python
t.test_2b(answer2_b)
```

**Question 2(c)** <br> {points: 1}  

Make the dataset tidy by using `.melt()`.    
Use `School` as `id_vars`, `Gender` as `var_name` and `Wage` as `value_name` and sort the dataframe by `School` and then `Gender`.      
Name the new object `tidy_wages`. 

<!-- grade_id: cell-73c51126a6542e2d, points: 0, grade: false, solution: true -->
```python
tidy_wages = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-1f14881b27d05b6a, points: 1, grade: true, solution: false -->
```python
t.test_2c(tidy_wages)
```

**Question 2(d)** <br> {points: 2}  

`.groupby()` the `Gender` column and find the mean `Wage` of each group.    
Save the new dataframe as `wage_means`. 

<!-- grade_id: cell-35d4b922fc450705, points: 0, grade: false, solution: true -->
```python
wage_means = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-c4891e7ec95d0a58, points: 2, grade: true, solution: false -->
```python
assert 'wage_means' in locals(), "Your answer should be named 'wage_means'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
```

**Question 2(e)** <br> {points: 1}  

Which of the following is True? 

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
 
A) Men and Women's mean wage are approximately equal. 

B) Women are making 0.9 the amount of Men  

C) Men are making 1.2 the amount of Women 

D) Women are making 0.7 the amount of Men 

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer2_e`.*

<!-- grade_id: cell-1ca40c4420470421, points: 0, grade: false, solution: true -->
```python
answer2_e = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-4db36636a741a9e1, points: 1, grade: true, solution: false -->
```python
t.test_2e(answer2_e)
```

**Question 2(f)** <br> {points: 1}  

Reverse the operation you performed in **Question 2(c)** using `.pivot()`. Do not forget to reset your index.
You will need to chain `.rename_axis('', axis='columns')` after you reset your axis to adjust for an incorrect axis label. 
Name the untidied object `untidied_wages`. 

<!-- grade_id: cell-d245d98a9ad389fa, points: 0, grade: false, solution: true -->
```python
# untidied_wages = (....pivot(...)
#                      ....()
#                      .rename_axis('', axis='columns')
#                  )
 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e89f064e31704eae, points: 1, grade: true, solution: false -->
```python
t.test_2f(untidied_wages)
```

**Question 2(g)** <br> {points: 1}  

Using `untidied_wages`  We are trying to answer the question, *Which school has the largest wage gap between among women and men?*

You will need to add a new column in the dataframe named `Wage_Gap` (Men wage - Woman wage). Sort your new dataframe in descending order and name it `wage_gap_df`.

<!-- grade_id: cell-2d8b2e33a5fda5a8, points: 0, grade: false, solution: true -->
```python
wage_gap_df = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

```python
wage_gap_df
```

<!-- grade_id: cell-d4a03759945efbad, points: 1, grade: true, solution: false -->
```python
t.test_2g(wage_gap_df)
```

**Question 2(h)** <br> {points: 1}  

Plot the dataframe `wage_gap_df` by specifying `School` on the `x-axis` and `Wage_Gap` on the `y-axis`.        
Name this object `wage_gap_plot`. Also sort the `Wage_Gap` values in ascending order. Don't forget to provide an informative title for your plot.

<!-- grade_id: cell-07402cfb0136e162, points: 0, grade: false, solution: true -->
```python
wage_gap_plot = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

wage_gap_plot
```

<!-- grade_id: cell-582c8f7c5ccfe601, points: 1, grade: true, solution: false -->
```python
t.test_2h(wage_gap_plot)
```

## 3.  Concat and Merge 

**Question 3(a)** <br> {points: 1}  

Combine the following 2 dataframes using either `pd.concat()` or `.merge()`. Drop any duplicate columns if any arise.    
Name the new dataframe `netflix_full`.


```python
netflix1 = pd.read_csv('data/netflix1.csv')
netflix1
```

```python
netflix2 = pd.read_csv('data/netflix2.csv')
netflix2
```

<!-- grade_id: cell-50875fc3cd80983e, points: 0, grade: false, solution: true -->
```python
netflix_full = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-74c7e427e77d492b, points: 1, grade: true, solution: false -->
```python
t.test_3a(netflix_full)
```

**Question 3(b)** <br> {points: 2}  

Add the new dataframe `new_releases` to the existing dataframe `netflix_full` that you made in the previous question.   
Name the new object `netflix_collection`.


```python
new_releases = pd.read_csv('data/netflix3.csv')
```

<!-- grade_id: cell-0fe1ae1c5d8700ab, points: 0, grade: false, solution: true -->
```python
netflix_collection = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
netflix_collection
```

<!-- grade_id: cell-015abf54b72e30df, points: 2, grade: true, solution: false -->
```python
assert 'netflix_collection' in locals(), "Your answer should be named 'netflix_collection'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
```

**Question 3(c)** <br> {points: 1} 

Here are 2 dataframes:


#### Left Dataframe


|      | Name                        | Gender   | Status     | Birth Date   | Birth Place     |
|---: |:----------------------------|:---------|:-----------|:-------------|:-----------------|
|**0**| Gregory E. Chamitoff        | Male     | Active     | 8/6/1962     | Montreal, Canada |
|**1**| Patricia Hilliard Robertson | Female   | Deceased   | 3/12/1963    | Indiana, PA      |
|**2**| Kathleen Rubins             | Female   | Active     | 10/14/1978   | Farmington, CT   |
|**3**| Michael T. Good             | Male     | Management | 10/13/1962   | Parma, OH        |
|**4**| Neil A. Armstrong           | Male     | Deceased   | 8/5/1930     | Wapakoneta, OH   |


#### Right  Dataframe


|     | Name       | Undergraduate Major  | Graduate Major   |   Space Flights |  Space Walks |   Space Walks (hr) |
|----:|:--------------------|:-------------------|:------------------|----------------:|--------:|-----------------:|
|**0**| Buzz Aldrin                 | Mechanical Engineering   | Astronautics          |     2 |     2 |    8 |
|**1**| Patricia Hilliard Robertson | Biology           | Medicine                     |     0 |     0 |    0 |
|**2**| Neil A. Armstrong           | Aeronautical Engineering | Aerospace Engineering |     2 |     1 |    2 |
|**3**| Mae C. Jemison              | Chemical Engineering     | Medicine              |     1 |     0 |    0 |
|**4**| Michael T. Good             | Aerospace Engineering    | Aerospace Engineering |     2 |     4 |   30 |
|**5**| Shannon W. Lucid            | Chemistry                | Biochemistry          |     5 |     0 |    0 |


The above dataframe have been joined together to produce each of the dataframes below, each one using a different joining methods. Match up each of the merged dataframe labelled A-D with it's corresponding join type ('inner', 'outer', 'left', 'right'). 


###### Dataframe A
|  | Name| Gender| Status| Birth Date| Birth Place| Undergraduate Major| Graduate Major|Space Flights |Space Walks |   Space Walks (hr) |
|---:|:----------------------------|:---------|:-----------|:-------------|:---------------|:-------------------------|:----------------------|----------------:|--------------:|-------------------:|
|**0**| Patricia Hilliard Robertson | Female| Deceased| 3/12/1963| Indiana, PA| Biology | Medicine  | 0 |0 |0 |
|**1**| Michael T. Good| Male|Management|10/13/1962|Parma, OH|Aerospace Engineering|Aerospace Engineering|2 |4|30 |
|**2**|Neil A. Armstrong|Male|Deceased| 8/5/1930|Wapakoneta, OH|Aeronautical Engineering|Aerospace Engineering|2|1|2|
|**3**| Buzz Aldrin | NaN | NaN | NaN  | NaN | Mechanical Engineering   | Astronautics |  2 |  2 | 8 |
|**4**| Mae C. Jemison  | NaN  | NaN | NaN  | NaN | Chemical Engineering | Medicine |  1 | 0 |  0 |
|**5**| Shannon W. Lucid   | NaN   | NaN  | NaN | NaN  | Chemistry | Biochemistry   | 5 |    0 |  0 |


###### Dataframe B

|  | Name| Gender| Status| Birth Date| Birth Place| Undergraduate Major| Graduate Major|Space Flights |Space Walks |   Space Walks (hr) |
|---:|:----------------------------|:---------|:-----------|:-------------|:---------------|:-------------------------|:----------------------|----------------:|--------------:|-------------------:|
|**0**| Gregory E. Chamitoff| Male| Active| 8/6/1962| Montreal, Canada | NaN| NaN | NaN |  NaN | NaN |
|**1**| Patricia Hilliard Robertson | Female| Deceased| 3/12/1963| Indiana, PA| Biology | Medicine  | 0 |0 |0 |
|**2**| Kathleen Rubins| Female | Active| 10/14/1978| Farmington, CT| NaN | NaN | NaN | NaN | NaN |
|**3**| Michael T. Good| Male|Management|10/13/1962|Parma, OH|Aerospace Engineering|Aerospace Engineering|2 |4|30 |
|**4**|Neil A. Armstrong|Male|Deceased| 8/5/1930|Wapakoneta, OH|Aeronautical Engineering|Aerospace Engineering|2|1|2|
|**5**| Buzz Aldrin | NaN | NaN | NaN  | NaN | Mechanical Engineering   | Astronautics |  2 |  2 | 8 |
|**6**| Mae C. Jemison  | NaN  | NaN | NaN  | NaN | Chemical Engineering | Medicine |  1 | 0 |  0 |
|**7**| Shannon W. Lucid   | NaN   | NaN  | NaN | NaN  | Chemistry | Biochemistry   | 5 |    0 |  0 |


###### Dataframe C

|  | Name| Gender| Status| Birth Date| Birth Place| Undergraduate Major| Graduate Major|Space Flights |Space Walks |   Space Walks (hr) |
|---:|:----------------------------|:---------|:-----------|:-------------|:---------------|:-------------------------|:----------------------|----------------:|--------------:|-------------------:|
|**0**| Patricia Hilliard Robertson | Female| Deceased| 3/12/1963| Indiana, PA| Biology| Medicine|0 |0 |0 |
|**1**|Michael T. Good| Male|Management|10/13/1962|Parma, OH|Aerospace Engineering|Aerospace Engineering|2 |4|30 |
|**2**|Neil A. Armstrong|Male|Deceased| 8/5/1930|Wapakoneta, OH|Aeronautical Engineering|Aerospace Engineering|2|1|2|

###### Dataframe D



|  | Name| Gender| Status| Birth Date| Birth Place| Undergraduate Major| Graduate Major|Space Flights |Space Walks |   Space Walks (hr) |
|---:|:----------------------------|:---------|:-----------|:-------------|:---------------|:-------------------------|:----------------------|----------------:|--------------:|-------------------:|
|**0**| Gregory E. Chamitoff| Male| Active| 8/6/1962| Montreal, Canada | NaN| NaN | NaN |  NaN | NaN |
|**1**| Patricia Hilliard Robertson | Female| Deceased| 3/12/1963| Indiana, PA| Biology | Medicine  | 0 |0 |0 |
|**2**| Kathleen Rubins| Female | Active| 10/14/1978| Farmington, CT| NaN | NaN | NaN | NaN | NaN |
|**3**| Michael T. Good| Male|Management|10/13/1962|Parma, OH|Aerospace Engineering|Aerospace Engineering|2 |4|30 |
|**4**|Neil A. Armstrong|Male|Deceased| 8/5/1930|Wapakoneta, OH|Aeronautical Engineering|Aerospace Engineering|2|1|2|



<!-- grade_id: cell-ed8bee364b3945b2, points: 0, grade: false, solution: true -->
```python
answer3c_dfa = None
answer3c_dfb = None
answer3c_dfc = None
answer3c_dfd = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-d4e2a89b6eb53de2, points: 1, grade: true, solution: false -->
```python
t.test_3c(answer3c_dfa, answer3c_dfb, answer3c_dfc, answer3c_dfd)
```

## 4. Green house emissions 

The file at this [url](https://raw.githubusercontent.com/plotly/datasets/master/Emissions%20Data.csv) contains data for the amount of greenhouse gas emissions in tonnes **per person** per year _(tons/person)_  for different countries around the world for the years 2008 - 2011. 

```python
emissions = pd.read_csv('https://raw.githubusercontent.com/plotly/datasets/master/Emissions%20Data.csv', dtype = {'Year':'object'})
```

##### Our mission
Our interest instead lies on finding the emission per year per person for each *continent* during the year 2010. It would not make sense to add up these current `Emission` variable as each country has a different population.  Instead,  we will need to find the raw quantity of gas emissions in tonnes per year for each country first, sum these per continent before dividing by the total continent population.

 **Question 4(a)** <br> {points: 1} 
 
 
To calculate a country's total green house gas emission from the per person value in the `Emission` column, we will need to obtain each country's population.  The `gapminder.csv` dataset has just that. It is a dataset located in the `/data` folder that has been minorly edited from the [Gapminder website](https://www.gapminder.org/data/) which is a tool and resource that promotes [sustainable global development and achievement of the United Nations Millennium Development Goals](https://www.gapminder.org/about-gapminder/our-mission/). It contains the population of each country from the years 1992 - 2020. Load only the `Country` column and the columns between 2008-2011 (You will have to put the column years in single quotations ex:`'2008'`) and name the dataframe `gapminder`. 
 

<!-- grade_id: cell-ed0d3d2ef1c6fe35, points: 0, grade: false, solution: true -->
```python
gapminder = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-77eea8f3cb8f5d48, points: 1, grade: true, solution: false -->
```python
t.test_4a(gapminder)
```

 **Question 4(b)** <br> {points: 1} 
 
Use `.melt()` to transform `gapminder` into a long dataframe. Name the new columns `Year` and `Population`.    
Name the new object `gapminder_long`.

<!-- grade_id: cell-1858fb5fc72c26c1, points: 0, grade: false, solution: true -->
```python
gapminder_long = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-622251d0f0a0ee83, points: 1, grade: true, solution: false -->
```python
t.test_4b(gapminder_long)
```

**Question 4(c)** <br> {points: 1} 
 
Use `.merge()` to join the dataframes `emissions` and `gapminder_long`. Use identifying columns `Country` and `Year` and use the `outer` join type. Make sure to add the `_merge` column with the `indicator` argument.    
Name the new object `emissions_pop`.


<!-- grade_id: cell-72b0faab1c274347, points: 0, grade: false, solution: true -->
```python
emissions_pop = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-af447271111330a6, points: 1, grade: true, solution: false -->
```python
t.test_4c(emissions_pop)
```

**Question 4(d)** <br> {points: 1} 

Let's put our concentration on just the year 2010. Filter so we only have the year 2010 and drop the column named `Year`.     
Name this dataframe `emissions_2010`. 

<!-- grade_id: cell-02da998eb73e152c, points: 0, grade: false, solution: true -->
```python
emissions_2010 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-69601f11edd342c0, points: 1, grade: true, solution: false -->
```python
t.test_4d(emissions_2010)
```

**Question 4(e)** <br> {points: 3} 

How many countries are missing from the left and right dataframes? Save the filtered shape (and using `[0]` to get the row dimension only) as objects named `missing_left`, `missing_right`. 

Create a new dataframe that only contains the rows that were present in both `emissions` and `gapminder_long` . 
Name this dataframe `emissions_both`. 

* Hint: Looking at the `_merge` column may be helpful.

<!-- grade_id: cell-f766c516446ac6bd, points: 0, grade: false, solution: true -->
```python
missing_left = None
missing_right = None
emissions_both = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-db9d5470fed86685, points: 3, grade: true, solution: false -->
```python
t.test_4e(missing_left, missing_right, emissions_both)
```

**Question 4(f)** <br> {points: 2} 

Make a new column for the total quantity of emissions for each country and name it `Quantity`. Currently we have the column `Emission` which is the greenhouse gas emissions in tonnes **per person** in the year 2010 and the number of people in the column `Population`. Use the `emissions_both` dataframe.

<!-- grade_id: cell-7fb38fc8cf061e42, points: 0, grade: false, solution: true -->
```python
emissions_total = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-6d405a90c9b81efa, points: 2, grade: true, solution: false -->
```python
assert 'emissions_total' in locals(), "Your answer should be named 'emissions_total'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
```

**Question 4(g)** <br> {points: 1} 

Group the `emissions_total` dataframe by continent and use `.agg()` to obtain the total sum of the `Population` and `Quantity` for each continent. Save this as `continent_emissions`. 

<!-- grade_id: cell-86efb0698414ae91, points: 0, grade: false, solution: true -->
```python
continent_emissions = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-cb25f5623b75069e, points: 1, grade: true, solution: false -->
```python
t.test_4g(continent_emissions)
```

**Question 4(h)** <br> {points: 1} 

Create a new column named `Emission_pp` from the `continent_emissions` dataframe which is the amount of emission **per person in each continent**. This can be made by dividing the `Quantity` column by the `Population` column.        
Save the dataframe as object `continent_emissions_pp`. 


<!-- grade_id: cell-c4a55a0ab6afbaef, points: 0, grade: false, solution: true -->
```python
continent_emissions_pp = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-99ac21b147484cd9, points: 1, grade: true, solution: false -->
```python
t.test_4h(continent_emissions_pp)
```

**Question 4(i)** <br> {points: 1} 

Plot in a bar graph, the `Emission_pp` for each continent from highest to lowest and name the plot `continent_emissions_pp_plot`. Give it a title as well.

Note that you will need to use the `reset_index` function.

<!-- grade_id: cell-3746a96f6e7c0c40, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-00c6e17519090236, points: 1, grade: true, solution: false -->
```python
t.test_4i(continent_emissions_pp_plot)
```

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.  

## Attributions
- Gapminder Dataset - [Gapminder](https://www.gapminder.org/data/)
- Astronaut Dataset - [Kaggle](https://www.kaggle.com/nasa/astronaut-yearbook?select=astronauts.csv)
- Netflix Dataset - [Kaggle](https://www.kaggle.com/shivamb/netflix-shows?select=netflix_titles.csv)
- MDS DSCI 523 - Data Wrangling course - [MDS's GitHub website](hhttps://ubc-mds.github.io/) 


## Module Debriefing

If this video is not showing up below, click on the cell and click the  button in the toolbar above.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('/X0eqYP2t3kA', width=854, height=480)
```

```python

```