## Introduction to Machine Learning  

## Assignment 7: Assessment and Measurements

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are the links of 2 videos that explain:

1. [How to save your Jupyter notebook](https://youtu.be/0aoLgBoAUSA) and,       
2. [How to answer a question in a Jupyter notebook assignment](https://youtu.be/7j0WKhI3W4s).

### Assignment Learning Goals:

By the end of the module, students are expected to:

- Explain why accuracy is not always the best metric in ML.
- Explain components of a confusion matrix.
- Define precision, recall, and f1-score and use them to evaluate different classifiers.
- Identify whether there is class imbalance and whether you need to deal with it.
- Explain `class_weight` and use it to deal with data imbalance.
- Appropriately select a scoring metric given a regression problem.
- Interpret and communicate the meanings of different scoring metrics on regression problems. MSE, RMSE, $R^2$, MAPE.
- Apply different scoring functions with `cross_validate` and `GridSearchCV` and `RandomizedSearchCV`.


This assignment covers [Module 7](https://ml-learn.mds.ubc.ca/en/module7) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` with your completed code and answers then proceed to run the cell!

Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct. These questions are worth 2 points.

```python
# Import libraries needed for this lab
from hashlib import sha1

import altair as alt
import graphviz
import numpy as np
import pandas as pd

from sklearn import tree
from sklearn.dummy import DummyClassifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.model_selection import train_test_split, cross_validate
from sklearn.neighbors import KNeighborsClassifier
from sklearn.impute import SimpleImputer
from sklearn.pipeline import Pipeline, make_pipeline
from sklearn.compose import make_column_transformer 
from sklearn.model_selection import GridSearchCV, RandomizedSearchCV
from sklearn.ensemble import RandomForestClassifier, RandomForestRegressor
from sklearn.preprocessing import (
    FunctionTransformer,
    Normalizer,
    OneHotEncoder,
    StandardScaler,
    normalize,
    scale)
from sklearn.metrics import make_scorer
from sklearn.svm import SVC, SVR

import test_assignment7 as t
#alt.renderers.enable('mimetype')
alt.data_transformers.disable_max_rows()
```

# 1. Precision, recall, and f1 score "by hand" (without `sklearn`)


Consider the problem of predicting whether a patient has a disease or not. Below are confusion matrices of two machine learning models: Model A and Model B. 

##### Model A
|    Actual/Predicted      | Predicted disease | Predicted no disease |
| :------------- | -----------------------: | -----------------------: |
| **Actual disease**       | 2 | 8 |
| **Actual no disease**       | 0 | 100 |


##### Model B
|    Actual/Predicted      | Predicted disease | Predicted no disease |
| :------------- | -----------------------: | -----------------------: |
| **Actual disease**       | 6 | 4 |
| **Actual no disease**       | 10 | 90 |

**Question 1.1** <br> {points: 1}  

Precision, recall, and f1 score depend crucially upon which class is considered "positive", that is the thing you wish to find. In the example above, which class is likely to be the "positive" class?

Save the label name in a string object named `answer_1_1`.

<!-- grade_id: cell-10bccd38560266fa, points: 0, grade: false, solution: true -->
```python
answer_1_1 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
answer_1_1
```

<!-- grade_id: cell-58c8cbee8ef77632, points: 1, grade: true, solution: false -->
```python
t.test_1_1(answer_1_1)
```

**Question 1.2** <br> {points: 3}  

Calculate accuracies for Model A and Model B. 

Save the values of each calculations as a fraction in objects name `model_a_acc` and `model_b_acc` respectively. 

<!-- grade_id: cell-da59f7cfdbd13642, points: 0, grade: false, solution: true -->
```python
model_a_acc = None
model_b_acc = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-112a1f1642f468fa, points: 1, grade: true, solution: false -->
```python
t.test_1_2_1(model_a_acc)
```

<!-- grade_id: cell-7adaec399eb35deb, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'model_b_acc' in globals(
), "Please make sure that your solution is named 'model_b_acc'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 1.3** <br> {points: 1}  

Which model would you pick simply based on the accuracy metric? 

Save either "Model A" or "Model B" in an object named `answer_1_3`.

<!-- grade_id: cell-88e8a579efb8b86f, points: 0, grade: false, solution: true -->
```python
answer_1_3 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-6d1051aef2f7edba, points: 1, grade: true, solution: false -->
```python
t.test_1_3(answer_1_3)
```

**Question 1.4** <br> {points: 3}  

Calculate precision, recall, f1-score for **Model A** by designating the appropriate fraction to objects named `a_precision`, `a_recall` and `a_f1`. 

You can use the objects `a_precision` and `a_recall` to use in your `a_f1` calculation.

<!-- grade_id: cell-4f35b1710839f096, points: 0, grade: false, solution: true -->
```python
a_precision = None
a_recall = None
a_f1 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-ebc2c04c38d967f8, points: 1, grade: true, solution: false -->
```python
t.test_1_4_1(a_precision)
```

<!-- grade_id: cell-e6873091826ecbfe, points: 1, grade: true, solution: false -->
```python
t.test_1_4_2(a_recall)
```

<!-- grade_id: cell-22acb42332686294, points: 1, grade: true, solution: false -->
```python
t.test_1_4_3(a_f1)
```

**Question 1.5** <br> {points: 3}  

Calculate precision, recall, f1-score for **Model B** by designating the appropriate fraction to objects named `b_precision`, `b_recall` and `b_f1`. 

You can use the objects `b_precision` and `b_recall` to use in your `b_f1` calculation.

<!-- grade_id: cell-1464d82d29d554be, points: 0, grade: false, solution: true -->
```python
b_precision = None
b_recall = None
b_f1 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-c41145a09c7bfc8b, points: 1, grade: true, solution: false -->
```python
t.test_1_5_1(b_precision)
```

<!-- grade_id: cell-6564d4b3d12c0f51, points: 1, grade: true, solution: false -->
```python
t.test_1_5_2(b_recall)
```

<!-- grade_id: cell-a97d0ee2aa39b97c, points: 1, grade: true, solution: false -->
```python
t.test_1_5_3(b_f1)
```

**Question 1.6** <br> {points: 1}  

Which metric is more informative in this case?

i) Accuracy

ii) Precision

iii) Recall

iv) f1



Select all that apply and add them into a list named `answer_1_6`. 
For example if statement i and iv are both true, your solution will look like this: 

```
answer_1_6 = ["i", "iv"] 
```


<!-- grade_id: cell-94a5d47d39c26a6e, points: 0, grade: false, solution: true -->
```python
answer_1_6 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-40689573c4b45ba7, points: 1, grade: true, solution: false -->
```python
t.test_1_6(answer_1_6)
```

**Question 1.7** <br> {points: 1}  

Which model would you pick based on this information? 

Save either "Model A" or "Model B" in an object named `answer_1_7`.

<!-- grade_id: cell-4dd94c25673682b2, points: 0, grade: false, solution: true -->
```python
answer_1_7 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-3b1fd3696243d2e6, points: 1, grade: true, solution: false -->
```python
t.test_1_7(answer_1_7)
```

# 2. Classification evaluation metrics using `sklearn`

In general, when a dataset is imbalanced, accuracy does not provide the whole story. In the module material, we looked at a credit card fraud dataset which is a classic example of an imbalanced dataset. Another example is customer churn datasets. For the next questions, you will be using a [bank customer churn dataset](https://www.kaggle.com/shubh0799/churn-modelling) from Kaggle. In this question, we will be concentrating on the target label `Exited`. 

```python
churn_df = pd.read_csv("data/churn.csv")
train_df, test_df = train_test_split(churn_df, test_size=0.3, random_state=123)
train_df.head()

```

**Question 2.1** <br> {points: 1}  

What is the distribution of target values (`Exited`) in the train split? Your answer should be of type Pandas Series and saved in an object named `class_dist`.

<!-- grade_id: cell-8fee174efb95046b, points: 0, grade: false, solution: true -->
```python
class_dist = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-9693a3487c29667f, points: 1, grade: true, solution: false -->
```python
t.test_2_1(class_dist)
```

**Question 2.2** <br> {points: 1}  

Let's now separate our feature vectors from the target.

Use all the columns except for `Exited` as your `X` and the `Exited` column as your target `y`. 

You will need to split both `train_df` and `test_df`. 

Save the results in objects named `X_train`, `y_train`, `X_test` and `y_test`. 

<!-- grade_id: cell-912921e7e74800fb, points: 0, grade: false, solution: true -->
```python
X_train, y_train, X_test, y_test = None, None, None, None
    
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

X_train.head()
```

<!-- grade_id: cell-f7c38b0ae6f70af2, points: 1, grade: true, solution: false -->
```python
t.test_2_2(X_train,X_test,y_train,y_test)
```

**Question 2.3** <br> {points: 1} 

Carry out cross-validation with `DummyClassifier` using the `stratified` strategy. Pass the following `scoring` metrics to `cross_validate`. 
- accuracy
- f1
- recall
- precision

Make sure you use  `return_train_score=True` and 5-fold cross-validation.

Save your results in a dataframe named `dummy_scores`. 

<!-- grade_id: cell-c6682256ddbc4a46, points: 0, grade: false, solution: true -->
```python
dummy_scores = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-f74782b5290c5565, points: 1, grade: true, solution: false -->
```python
t.test_2_3(dummy_scores)
```

**Question 2.4** <br> {points: 1} 

What is the mean of each column in `dummy_scores`?


Save your result in an object named `dummy_mean`. 

<!-- grade_id: cell-99de59fe13557284, points: 0, grade: false, solution: true -->
```python
dummy_mean = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-57af51df99cb0bdf, points: 1, grade: true, solution: false -->
```python
t.test_2_4(dummy_mean)
```

**Question 2.5** <br> {points: 5}  

Using either `.describe(include="all")`, `.info()` or looking at the data, split it into 4 types of features; numeric, categorical, binary and `drop_features`. 

- Add the labels of the numeric column(s) (as type string) to a list name `numeric_features`.
- Add the labels of the categorical column(s) (as type string) to a list name `categorical_features`.
- Add the labels of the binary column(s) (as type string) to a list name `binary_features`.
- Add the labels of the column that should be excluded from the model fitting to a list named `drop_features`. _(Hint: which of the columns is a unique identifier for the examples?)_

Identify different feature types (e.g., numeric, categorical, binary, drop features).
train_df.info()

```python
X_train.columns
```

<!-- grade_id: cell-66dbf242bf1a9faa, points: 0, grade: false, solution: true -->
```python
numeric_features = None
categorical_features = None
drop_features = None
binary_features = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b422078fdd00a9a8, points: 1, grade: true, solution: false -->
```python
t.test_2_5_1(numeric_features)
```

<!-- grade_id: cell-fd94cbdf40c54387, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'categorical_features' in globals(
), "Please make sure that your solution is named 'categorical_features'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

<!-- grade_id: cell-9260945a70a8df44, points: 1, grade: true, solution: false -->
```python
t.test_2_5_3(drop_features)
```

<!-- grade_id: cell-8a7d80a0f3f69e0b, points: 1, grade: true, solution: false -->
```python
t.test_2_5_4(binary_features)
```

Since we have no null values we can skip the individual pipelines and make our column transforming with just a single transformation. 

Because everyone needs a little help once in a while, we are going to do this part for you! 

```python
preprocessor = make_column_transformer(
    ("drop", drop_features),
    (StandardScaler(), numeric_features),
    (OneHotEncoder(handle_unknown="ignore"), categorical_features),
    (OneHotEncoder(handle_unknown="error", drop="if_binary"), binary_features)
)
```

**Question 2.6** <br> {points: 2}  

In this question, you will be using the [`RandomForestClassifier`](https://scikit-learn.org/stable/modules/generated/sklearn.ensemble.RandomForestClassifier.html) which you haven't studied yet but did see in assignment 3.  You should feel comfortable using models with our usual ML workflow even if you haven't seen them before. 


Build a pipeline named `unbalanced_pipe` that first preprocesses `preprocessor` and then builds a  `RandomForestClassifier` using a random_state of 77. 

Carry out cross-validation on `unbalanced_pipe` and the training set using the `cross_validate` function and the following evaluation metrics:
- `accuracy`
- `precision`
- `recall`
- `f1`

Note that you can pass multiple [scoring metrics](https://scikit-learn.org/stable/modules/model_evaluation.html#scoring-parameter) as a list or a dict to the `scoring` parameter. 

Save the results of the cross-validation in a dataframe named`rf_unbalanced_scores`.

<!-- grade_id: cell-0de662581b33ae12, points: 0, grade: false, solution: true -->
```python
unbalanced_pipe = None
rf_unbalanced_scores = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e578cb24650f0946, points: 1, grade: true, solution: false -->
```python
t.test_2_6(rf_unbalanced_scores)
```

**Question 2.7** <br> {points: 1} 

What is the mean of each column in `rf_unbalanced_scores`?


Save your result in an object named `rf_unbalanced_mean`. 

<!-- grade_id: cell-22fba680c83299d9, points: 0, grade: false, solution: true -->
```python
rf_unbalanced_mean = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-472ac701d5761924, points: 1, grade: true, solution: false -->
```python
t.test_2_7(rf_unbalanced_mean)
```

**Question 2.8** <br> {points: 2}  

Repeat question 6 above but this time set `class_weight="balanced"` in the `RandomForestClassifier`. 
Save the new pipeline in an object named `balanced_pipe`. Don't forget to use `random_state=77`for the classifier. 

Carry out cross-validation on `balanced_pipe` but this time save the scores in a dataframe named `rf_balanced_scores`.

<!-- grade_id: cell-1c4687b4a9a6aac8, points: 0, grade: false, solution: true -->
```python
balanced_pipe = None
rf_balanced_scores = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-49f5395fd9aa9edd, points: 2, grade: true, solution: false -->
```python
t.test_2_8(rf_balanced_scores)
```

**Question 2.9** <br> {points: 1} 

What is the mean of each column in `rf_balanced_scores`?


Save your result in an object named `rf_balanced_mean`. 

<!-- grade_id: cell-1d3c898d4c72428c, points: 0, grade: false, solution: true -->
```python
rf_balanced_mean = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-f57fbaf001a3bd23, points: 1, grade: true, solution: false -->
```python
t.test_2_9(rf_balanced_mean)
```

**Question 2.10** <br> {points: 1} 

Look at your results above. Which of the following statements are true? 

i) Both random forest models have better mean validation accuracy performance than the dummy classifier.

ii) The balanced class random forest model has a better mean validation recall score than the unbalanced random forest model. 

iii) The balanced class random forest model has a better mean validation precision score than the unbalanced random forest model.  

iv) Mean validation precision and recall scores only increased in the balanced class random forest model.

v) Both random forest models show higher validation precision scores than validation recall scores. 


Select all that apply and add them into a list named `answer_2_10`. 
For example if statement i and iv are both true, your solution will look like this: 

```
answer_2_10 = ["i", "iv"] 
```


<!-- grade_id: cell-6de147c823c76215, points: 0, grade: false, solution: true -->
```python
answer_2_10 = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-fd3c4b417ac834cf, points: 1, grade: true, solution: false -->
```python
t.test_2_10(answer_2_10)
```

**Question 2.11** <br> {points: 1} 

For this next question, we have given you the majority of the code to hyperparameter tune the balanced `RandomForestClassifier`. 

Copy and paste the contents of the cell below into the following cell where the solution should be. Fill in the `....` blanks below so that it executes and returns the output. 

<!-- grade_id: cell-02e169611793d840, points: 0, grade: false, solution: true -->
```python
import scipy
from scipy.stats import randint

rf_pipeline = make_pipeline(
    preprocessor, ....(class_weight=...., random_state=123)
)

param_dist = {
    "randomforestclassifier__n_estimators": scipy.stats.randint(low=10, high=300),
    "randomforestclassifier__max_depth": scipy.stats.randint(low=2, high=20),
}

random_search = ....(
    ....,
    ....,
    n_iter=50,
    cv=3,
    verbose=1,
    n_jobs=-1,
    ....="f1",
    random_state=123,
)
random_search.....(X_train, y_train)

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-329ba0ee56f73721, points: 1, grade: true, solution: false -->
```python
t.test_2_11(rf_pipeline, random_search)
```

**Question 2.12** <br> {points: 3}

What are the best hyperparameter value for `n_estimators` and `max_depth`. Save it in an object named `optimal_parameters` (The auto-grader is expecting a dictionary object). 

What was the corresponding validation score? Save this in an object named `optimal_score`. 

*Hint: `.best_params_`  and `.best_score_` are helpful here.* 


<!-- grade_id: cell-e5054e9f65a2ecdc, points: 0, grade: false, solution: true -->
```python
optimal_parameters = None
optimal_score = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-f6d885ec6d0c2d90, points: 1, grade: true, solution: false -->
```python
t. test_2_12_1(random_search, optimal_parameters)
```

<!-- grade_id: cell-1f755ca4548b90ca, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'optimal_score' in globals(
), "Please make sure that your solution is named 'optimal_score'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 2.13** <br> {points: 2} 

What is the train and test score of the best scoring model? Save the result in objects named `training_score`, and `testing_score` respectively.

<!-- grade_id: cell-f0ab8c39cbe9ed7b, points: 0, grade: false, solution: true -->
```python
training_score = None
testing_score = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b8e8776e62548a4e, points: 1, grade: true, solution: false -->
```python
t.test_2_13_1(training_score)
```

<!-- grade_id: cell-61c66a15b802b184, points: 1, grade: true, solution: false -->
```python
t.test_2_13_2(testing_score)
```

**Question 2.14** <br> {points: 1} 

Import the appropriate libraries to plot a confusion matrix and print a classification report. 

<!-- grade_id: cell-1329e2207ccdb878, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-7cef3fe5c8c08343, points: 1, grade: true, solution: false -->
```python
t.test_2_14()
```

**Question 2.15** <br> {points: 1} 


Plot a confusion matrix on the test set using the object `random_search` as your estimator and "normalize" all your results (see the [documentation](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.plot_confusion_matrix.html) for more help here).

Name the plot `cm_plot`. 

<!-- grade_id: cell-d4ee02e8d0ad2f1f, points: 0, grade: false, solution: true -->
```python
cm_plot = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e926b5a0a1f975c5, points: 1, grade: true, solution: false -->
```python
t.test_2_15(cm_plot)
```

**Question 2.16** <br> {points: 3} 

Below print a classification report on the `X_test` predictions of `random_search`'s best model. Use this information to answer the questions below.

```python
# Use this cell to print your classification report
```



A) What is the recall considering that `Exited`(=1) is our "positive" class? Save the result to two decimal places in an object named `answer2_16a`. 

B) What is the precision weighted average? Save the result to two decimal places in an object named `answer2_16b`. 

C) How many customers exited in this test set? Save the result as an integer in an object named `answer2_16c`.

<!-- grade_id: cell-2a0327eb7dbf23a9, points: 0, grade: false, solution: true -->
```python
answer2_16a = None
answer2_16b = None
answer2_16c = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b4b39de2552f60bc, points: 1, grade: true, solution: false -->
```python
t.test_2_16_1(answer2_16a)
```

<!-- grade_id: cell-c4325891305a9cae, points: 1, grade: true, solution: false -->
```python
t.test_2_16_2(answer2_16b)
```

<!-- grade_id: cell-13dafcb469b21ef5, points: 1, grade: true, solution: false -->
```python
t.test_2_16_3(answer2_16c)
```

**Question 2.17** <br> {points: 1} 

What's happened to our model since we have tuned our hyperparameters?

A) We have sacrificed some of our precision score for a better recall score.

B) Our accuracy has decreased.

C) We have likely increased the number of falsely identified exited customers.

D) All of the above.

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer2_17`.*

<!-- grade_id: cell-edf2e1be7eceed82, points: 0, grade: false, solution: true -->
```python
answer2_17 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
answer2_17
```

<!-- grade_id: cell-f0db9a56f75d6b9a, points: 1, grade: true, solution: false -->
```python
t.test_2_17(answer2_17)
```

# 3. Regression Metrics

For this exercise, we will bring back the [Taiwan real estate valuation dataset](https://archive.ics.uci.edu/ml/datasets/Real+estate+valuation+data+set) that we saw in assignment 1. 

We are doing the first part for you. 

```python
housing_df = pd.read_csv('data/real_estate.csv')
train_df, test_df = train_test_split(housing_df, test_size=0.2, random_state=123)
train_df.head()
```

```python
X_train = train_df.drop(columns=['price'])
y_train = train_df['price']

X_test = test_df.drop(columns=['price'])
y_test = test_df['price']
```

**Question 3.1** <br> {points: 1}  

Since we have seen this dataset already, we know there are only numeric features and no missing values so we can skip quite a few steps here. 

Make a single pipeline using `make_pipeline` named `reg_pipe` that first uses `StandardScaler()`.


<!-- grade_id: cell-69d4e91a504c1837, points: 0, grade: false, solution: true -->
```python
reg_pipe = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-d8d5d76226e3a145, points: 1, grade: true, solution: false -->
```python
t.test_3_1(reg_pipe)
```

**Question 3.2** <br> {points: 1}  

Given the MAPE function we have given below, create a scorer function that we can pass into cross-validation. (look [here](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.make_scorer.html?highlight=make_scorer#sklearn.metrics.make_scorer) for help). 

Save it in an object named `mape_scorer`. Remember that MAPE is an error measure so we must use `greater_is_better=False` as an argument. 


<!-- grade_id: cell-bfca9f71be42499d, points: 0, grade: false, solution: true -->
```python
def mape(true, pred):
    return 100.0 * np.mean(np.abs((pred - true) / true))


# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-8a3cf905241ca8b0, points: 1, grade: true, solution: false -->
```python
t.test_3_2(mape_scorer)
```

**Question 3.3** <br> {points: 1}  

Make a dictionary named `scoring_dict` consisting the following metrics: 

- `neg_mean_squared_error`
- `neg_root_mean_square`
- `neg_mean_absolute_error`
- `r2`
- `mape_scorer`

The key-value pairs in this dictionary should be identical. (Aka the item to the left of the colon is the same as the right side of the colon - we have started it for you just in case.)

Remember that `mape_scorer` is a variable and should **NOT** be called as a string. (Although we need to name it as a string value for the dictionary's key; See slide 6 of exercise 21 for help.)

<!-- grade_id: cell-b824fb2b08c194e2, points: 0, grade: false, solution: true -->
```python
scoring_dict = {"neg_mean_squared_error": "neg_mean_squared_error",
               }

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-2262d2645f47df97, points: 1, grade: true, solution: false -->
```python
t.test_3_3(scoring_dict)
```

**Question 3.4** <br> {points: 1}  

Carry out cross-validation on `reg_pipe` and the training set using the `cross_validate` function.

Make sure to set `return_train_score=True` and assign our `scoring_dict` to the `scoring` argument.


Save the results of the cross-validation in a dataframe named `regression_scores`.

<!-- grade_id: cell-457e7772f822ea07, points: 0, grade: false, solution: true -->
```python
regression_scores = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-db06848393467bff, points: 1, grade: true, solution: false -->
```python
t.test_3_4(regression_scores)
```

**Question 3.5** <br> {points: 1} 

What is the mean of each column in `regression_scores`?


Save your result in an object named `regression_mean`. 

<!-- grade_id: cell-cfb55ddae37d8d25, points: 0, grade: false, solution: true -->
```python
regression_mean = None 

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-4993cba10fcfe8c5, points: 1, grade: true, solution: false -->
```python
t.test_3_5(regression_mean)
```

**Question 3.6** <br> {points: 2} 

What is the model's mean validation MAPE score? 

Save your answer to 2 decimal places in an object named `mean_reg_mape`. 

<!-- grade_id: cell-9cfd9da29de1973e, points: 0, grade: false, solution: true -->
```python
mean_reg_mape = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-763f0a302ba20680, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'mean_reg_mape' in globals(
), "Please make sure that your solution is named 'mean_reg_mape'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

## Attributions
- The Bank Customer Churn DataSet - [Kaggle](https://www.kaggle.com/shubh0799/churn-modelling)

- Real Estate Dataset - [The UCI Machine Learning Repository](https://archive.ics.uci.edu/ml/datasets/Real+estate+valuation+data+set)

*Yeh, I. C., & Hsu, T. K. (2018). Building real estate valuation models with comparative approach through case-based reasoning. Applied Soft Computing, 65, 260-271.*


- MDS DSCI 573 - Feature & Model Selection  - [MDS's GitHub website](https://github.com/UBC-MDS/DSCI_573_feat-model-select) 


## Before Submitting 

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.  