# Programming in Python for Data Science 

# Assignment 5: Making Choices and Repeating Iterations

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are the links of 2 videos that explain:

1. [How to save your Jupyter notebook](https://youtu.be/0aoLgBoAUSA) and,       
2. [How to answer a question in a Jupyter notebook assignment](https://youtu.be/7j0WKhI3W4s).  

### Assignment Learning Goals:

By the end of the module, students are expected to:

- Explain the DRY principle and how it can be useful.
- Write conditional statements with `if`, `elif` and `else` to run different code, depending on the input.
- Write `for` loops to repeatedly run code.
- Describe the expected outcome of code with nested loops.
- Define and use a function that accepts parameters and returns values.


This assignment covers [Module 5](https://prog-learn.mds.ubc.ca/en/module5) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` and the `raise NotImplementedError # No Answer - remove if you provide an answer` with your completed code and answers then proceed to run the cell!

Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct. These questions are worth 2 points.

```python
# Import libraries needed for this lab
import ast
from hashlib import sha1
import pandas as pd
from random import random
import test_assignment5 as t
```

## 1.  Conditionals 

**Question 1(a)** <br> {points: 1} 

The following lines of code have been scrambled into the wrong order. Uncomment the code below into your answer cell and rearrange the lines below so that it no longer results in an error. To help you, here's the logic behind the correct code in written English:

- if the temperature is less than 5, clothing should be a Heavier Coat
- if the temperature is between 5 (including) and 10 (excluding), clothing should be a Light Jacket
- if the temperature is 10 degrees or more, clothing should be a Light Sweater

<!-- grade_id: cell-025b7dd0f3c0f587, points: 0, grade: false, solution: true -->
```python
# clothing
# else: 
# temperature = 8
# elif temperature < 10:
#     clothing = 'Heavier Coat'
#     clothing = 'Light Jacket'
# if temperature < 5:
#     clothing = "Light Sweater"

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-dface204a341b0b2, points: 1, grade: true, solution: false -->
```python
t.test_1a(clothing)
```

**Question 1(b)** <br> {points: 1} 
 

Write conditional statements that categorizes a diamond's `size` by the following requirements: 

- Any size including 0.5 carats or less is a "Small diamond" 

- Anything between 0.5 (excluding) and 1.5 (including) carats is an "Average-sized diamond"

- Anything between 1.5 (excluding) and 2.5 (including) carats is a "Large diamond" and,

- Any size greater than 2.5 (excluding) carats is a "Huge diamond"


Start with an object named `diamond` with a value of 1.5. 

<!-- grade_id: cell-44485888f5604a95, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-a0183b67993f8245, points: 1, grade: true, solution: false -->
```python
t.test_1b(size)
```

**Question 1(c)** <br> {points: 1} 

Hussein categorizes his pantry type as either vegan, vegetarian or neither based on the conditional statements below:
```
if 'milk' not in pantry and 'meat' not in pantry:
    pantry_type = "vegan"
elif 'meat' not in pantry: 
    pantry_type = "vegetarian"
else:
    pantry_type = "neither"
```

       
Given `pantry` below, what is Hussein's `pantry_type`? Save the correct answer in a variable named `answer1_c`.


```python
pantry = {'pasta': 3, 'garlic': 4,'sauce': 2,
          'basil': 2, 'salt': 3, 'olive oil': 3,
          'rice': 3, 'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 1, 'onions': 1, 'mushrooms': 3,
          'broccoli': 2, 'butter': 2,'pickles': 6, 'milk': 2,
          'chia seeds': 5}
```

<!-- grade_id: cell-1446d930bb638d95, points: 0, grade: false, solution: true -->
```python
answer1_c = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-946e99977890301b, points: 1, grade: true, solution: false -->
```python
t.test_1c(answer1_c)
```

**Question 1(d)** <br> {points: 1} 

Hussein's friends like his categorizing trick, so they try to replicate it but are not as successful. 
Frida's code looks like this: 
   

```python
if 'milk' not in pantry and 'meat' not in pantry:
    pantry_type = "vegan"
if 'meat' not in pantry:  
    pantry_type = "vegetarian"
else:
    pantry_type = "neither"
```


Which of the following scenarios is likely to occur?


A) If the pantry contains meat and milk, the pantry will be categorized as "vegetarian" because of the second if statement.

B)  If the pantry contains milk and not meat, the pantry will be categorized as "neither" because the conditionals are  only and if-else statement.


C)  If the pantry doesn't contain milk or meat, the pantry will be categorized as both "vegetarian" and "neither" because of the second if statement.


D)  If the pantry doesn't contain milk and meat, the pantry will be categorized as both "vegan" and "vegetarian" because of the second if statement.


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between "", assign the correct answer to an object called `answer1_d`.*

<!-- grade_id: cell-c7b2a3e2d5fdee4a, points: 0, grade: false, solution: true -->
```python
answer1_d = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-139ebfa7a4f09d4e, points: 1, grade: true, solution: false -->
```python
t.test_1d(answer1_d)
```

**Question 1(e)** <br> {points: 2} 

Another one of Hussein friends likes his categorizing trick too, so they also try to replicate it but again they don't do it correctly. 
Their code looks like this: 
 

```python
if 'milk' and 'meat' not in pantry: 
    pantry_type = "vegan"
elif 'meat' not in pantry: 
    pantry_type = "vegetarian"
else:
    pantry_type = "neither"
```



Which of the following scenarios is likely to occur?


A) We learned in Module 4 that non-empty strings get cast to true. Since the first condition is just `if milk` and is missing `not in pantry`, this portion of the `if` condition casts always to true.  If the pantry contains milk and no meat, the pantry will be incorrectly typed as "vegan".

B) Because it is missing the `not in pantry` part of the first condition, the code won't work and an error will occur. 

C) We learned in Module 4 that non-empty strings get cast to true. Since the first condition is just `if milk` and is missing the `not in pantry`, this portion of the `if` condition casts always to true.  If the pantry does **not** contain milk or meat, the pantry will be incorrectly typed as "vegan".

D) The missing `not in pantry` causes this `if` statement to always equal true, so any combination of ingredients will result in a "vegan" pantry type. 


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between "", assign the correct answer to an object called `answer1_e`.*

<!-- grade_id: cell-ec4116fb89d0f246, points: 0, grade: false, solution: true -->
```python
answer1_e = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b91b12d96008b503, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer1_e' in globals(
), "Please make sure that your solution is named 'answer1_e'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 1(f)** <br> {points: 1} 

Frida tries to fix her code but this time she's changed it so the code looks like this: 
    
  

```python
if 'meat' not in pantry:  
    pantry_type = "vegetarian"
elif 'milk' not in pantry and 'meat' not in pantry:
    pantry_type = "vegan"
else:
    pantry_type = "neither"
```

Which of the following scenarios is likely to occur?

A) If meat and milk are both present, the pantry will be incorrectly categorized as "vegan".

B) If milk is present in the pantry but no meat, the pantry will be incorrectly categorized as "vegetarian".

C) If neither milk or no meat are present, the pantry will be incorrectly categorized as "vegetarian".

D) If no meat is present, the pantry will be categorized as "neither".


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between "", assign the correct answer to an object called `answer1_f`.*

<!-- grade_id: cell-f9d7fbab0ddb4359, points: 0, grade: false, solution: true -->
```python
answer1_f = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-3348cf47899adf79, points: 1, grade: true, solution: false -->
```python
t.test_1f(answer1_f)
```

**Question 1(g)** <br> {points: 1} 

Hussein wants to make a meal and have some friends over for a dinner party. He needs to know how complex his meal is going to be based on which ingredients he has in his pantry. 


In the cell below, you will see code that has been commented out (`#`). Correct the indentation for each if and else statement in the code below so that it run without errors ***AND*** returns the correct answer based on the ingredients needed.


Given Hussein's pantry from **question 1 (c)**, What does `party` return? 
Answer by saving the result in an object named `dinner`.  

**HINT:** You do not need to arrange the lines, you just need to change the indentation.

<!-- grade_id: cell-2355a6be939ed14a, points: 0, grade: false, solution: true -->
```python
pantry = {'pasta': 3, 'garlic': 4,'sauce': 2,
          'basil': 2, 'salt': 3, 'olive oil': 3,
          'rice': 3, 'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 1, 'onions': 1, 'mushrooms': 3,
          'broccoli': 2, 'butter': 2,'pickles': 6, 'milk': 2,
          'chia seeds': 5}

dinner = None

# if 'pasta' in pantry and 'sauce' in pantry and 'basil' in pantry:
# if pantry['pasta'] > 2:
# if pantry['sauce'] > 2:
# if pantry['basil'] > 3:
# party = "Basil and Tomato Pasta for 4 guests "
# else: 
# party = "Marinara Pasta for 4 Guests"
# else: 
# party = "Tomato Pasta with Basil Garnish for 2 guests"
# else: 
# party = "Tomato Pasta for 1 guest"
# else:
# party = "He can't make this recipe" 


# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-b719b82a30b0cec3, points: 1, grade: true, solution: false -->
```python
t.test_1g(dinner)
```

## 2. Loops

**Question 2(a)** <br> {points: 1} 

Write a `for` loop that produces the following output and saves each word as an element in a list named `new_list`.:

    For
    every
    word
    in
    this
    sentence,
    print
    the
    word
    on
    its
    own
    line
_Hint: .split() and print() may come in handy here._

<!-- grade_id: cell-c65f36e248846565, points: 0, grade: false, solution: true -->
```python
sentence = "For every word in this sentence, print the word on its own line"
new_list = []
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-453e5a996c1cb79d, points: 1, grade: true, solution: false -->
```python
t.test_2a(new_list)
```

**Question 2(b)** <br> {points: 1} 

Hussein wants to go to the grocery store to pick up some supplies for his dinner party. He wants to know how much stock he has for each ingredient so he doesn't buy too many groceries.    

Create a loop to obtain the values from the `pantry` dictionary and append them to the list name `ingredient_stock`

_Hint: You only need one loop here._


<!-- grade_id: cell-ebb671992e967ddb, points: 0, grade: false, solution: true -->
```python
pantry = {'pasta': 3, 'garlic': 4,'sauce': 2,
          'basil': 2, 'salt': 3, 'olive oil': 3,
          'rice': 3, 'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 1, 'onions': 1, 'mushrooms': 3,
          'broccoli': 2, 'butter': 2,'pickles': 6, 'milk': 2,
          'chia seeds': 5}

ingredient_stock = list()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-c552fbee27a15b2a, points: 1, grade: true, solution: false -->
```python
t.test_2b(ingredient_stock)
```

**Question 2(c)** <br> {points: 2} 

What is the total quantity of ingredients in his pantry? Save this in an object named `stock_total`.

_Hint: No loop is necessary here._

<!-- grade_id: cell-0b94c62481cc63bf, points: 0, grade: false, solution: true -->
```python
stock_total = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-5cc8eb93474daa65, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'stock_total' in globals(
), "Please make sure that your solution is named 'stock_total'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 2(d)** <br> {points: 1} 

Hussein is planning ahead. In case he loses his list, he wants to know the mean ingredient quantity before he leaves.    
Calculate is the mean storage space of Hussein's ingredients and save this value in an object named `stock_mean`.

Hint: You can use the object `ingredient_stock` from the previous question.

<!-- grade_id: cell-f0e86af08a7df00c, points: 0, grade: false, solution: true -->
```python
stock_mean = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-6071a5eea5217a2b, points: 1, grade: true, solution: false -->
```python
t.test_2d(stock_mean)
```

**Question 2(e)** <br> {points: 2} 

Hussein doesn't want to stop at just the mean, he wants as much information as possible now. He decided he also wants to calculate the pantry's ingredient quantity variance.

**sample variance** can be calculated from the sum of squared differences of each observation from the mean:    

$variance = \frac{\Sigma{(x-mean)^2}}{n}$

where **mean** is the mean of our observations, **x** is each individual observation, and **n** is the number of observations.

For this question:
   
- Create a loop that iterates over the `ingredient_stock` where for each ingrediant, you subtract the mean and square the result. (The mean is the number you calculated in **Question 2(d)** and x is each element in the list `ingredient_stock`)

- Add this value to `sum_diffsq` at each iteration.
  
Note:

- *Not every calculation has to be within the loop.* Once you have summed up all the squared values, you can take the `sum_diffsq` and divide it by the number of observations in the list `ingredient_stock`. 
   
- *To take a number to a power in Python, use `**`, e.g. `5 ** 2` for 5 squared.*


<!-- grade_id: cell-8126c9dd3dcbf70a, points: 0, grade: false, solution: true -->
```python
sum_diffsq = 0
stock_variance= None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-afbd5bead54163bc, points: 2, grade: true, solution: false -->
```python
t.test_2e(stock_variance)
```

**Question 2(f)** <br> {points: 2} 

Hussein is making spaghetti tonight and he needs to make sure that if he doesn't have enough of the ingredients in his pantry, he adds them to his shopping list. 

- For each item in the recipe, check if the ingredient is in Hussein's pantry. 
- If the recipe ingredient is in the pantry, check if the recipe requires more of the ingredient than what Hussein has in his pantry. If so, add the name and the quantity Hussein needs to purchase as key-value pairs in the dictionary `shopping_list`. 
- If the recipe item is not in the pantry, add the ingredient and the quantity as key-value pairs in the dictionary `shopping_list`.

<!-- grade_id: cell-c677a207ed1b7aa2, points: 0, grade: false, solution: true -->
```python
pantry = {'pasta': 3, 'garlic': 4,'sauce': 2,
          'basil': 2, 'salt': 3, 'olive oil': 3,
          'rice': 3, 'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 1, 'onions': 1, 'mushrooms': 3,
          'broccoli': 2, 'butter': 2,'pickles': 6, 'milk': 2,
          'chia seeds': 5}

meal_recipe = {'pasta': 2, 'garlic': 2, 'sauce': 3,
          'basil': 4, 'salt': 1, 'pepper': 2,
          'olive oil': 2, 'onions': 2, 'mushrooms': 6}


shopping_list = dict()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer

```

<!-- grade_id: cell-f65167d39d3ce5c2, points: 2, grade: true, solution: false -->
```python
t.test_2f(shopping_list)
```

## 3. Implementing the DRY Principle

**Question 3(a)** <br> {points: 1} 

In the last 4 days, Arzan's average pulse for his heart rate has been the following beats per minute: 72, 88, 94, 65. He wants Python to decide which category each day's heart rate belongs to. The possible categories are:
- "High heart rate" if it's greater than 80
- "Average  heart rate" for anything between 70-80, and
- "Low  heart rate" for anything below 70.

Similar to what we did in the first section of the Module's learning material, **without using a loop**, make conditional statements for each day's pulse so Python can add each pulse to a list named `heartrate_labels`. 

*Hint: You should have an `if`, `elif` and `else` statement for **each** day*

<!-- grade_id: cell-0d8da57d062c2642, points: 0, grade: false, solution: true -->
```python
pulse_day1 = 72
pulse_day2 = 88
pulse_day3 = 94
pulse_day4 = 65

heartrate_labels = list()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-1e74868c09310ad1, points: 1, grade: true, solution: false -->
```python
t.test_3a(heartrate_labels)
```

```python
heartrate_labels
```

**Question 3(b)** <br> {points: 1} 

Repeat the question above, but this time use a `for` loop.

<!-- grade_id: cell-4b834c41897dd8d9, points: 0, grade: false, solution: true -->
```python
pulses = [pulse_day1, pulse_day2, pulse_day3, pulse_day4]

heartrate_labels = list()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-844eb90934065472, points: 1, grade: true, solution: false -->
```python
t.test_3b(heartrate_labels)
```

**Question 3(c)** <br> {points: 2} 

Which of the following are advantages of the code implemented in Question 3(b) vs Question 3(a)?

A) Changing the high pulse threshold from 80 to 90 would only require one change in the code instead of 4 changes.

B) The code works for a list of any length.

C) The results will be more correct.

D) The code takes less time to read.

E) The code adheres to the DRY principle.


*Assign your answer to an object called `answer2_i` in the code chunk below. Make sure your answer contains uppercase letters and surround it with quotation marks and square brackets.  If there are more than one answer to this question, separate each letter with a comma within the square brackets. For example, if you believe the answer is A, B and C your answer would like this:   
`answer2_i = ['A', 'B', 'C')`*

<!-- grade_id: cell-dbc77eeaf1e0d005, points: 0, grade: false, solution: true -->
```python
answer3_c = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-bc861bd9ad97766d, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'answer3_c' in globals(
), "Please make sure that your solution is named 'answer3_c'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

## 4. Nested Loops

**Question 4(a)** <br> {points: 1} 

Given the following code below, what is the output?


<img src='static/loop0.png' width="560">

Answer in the cell below and assign the correct answer to an object called count.

<!-- grade_id: cell-a6198f50e4ef53b3, points: 0, grade: false, solution: true -->
```python
count = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-76615b209463f859, points: 1, grade: true, solution: false -->
```python
t.test_4a(count)
```

**Question 4(b)** <br> {points: 3} 

Hussein is at the store and some of the ingredients on his shopping list are organized by department.    
```
ingredient_categories = {'dairy':['butter', 'milk', 'cheese'], 'produce':['potatoes', 'onion', 'broccoli', 'mushrooms'], 'spices': ['salt', 'pepper', 'paprika']}
```
Hussein wants to know how many items from his pantries he has in each of the ingredient_categories `dairy`, `produce` and `spices`. We have given you the outer part of the code to help start you off. 

This question is a little trickier so use `print()` to check what is being output in the loops and variables. 


_Hint:_

- *Create an inner loop that iterates over all the items in the `ingredients`.*
- *Make sure you include a conditional statement that only compares ingredients that are in Hussein's pantry.*
- *If it passes the conditional statement, add the pantry's value to the `category`'s existing value in `category_counts` so it sums up the total quantity of ingredients for each of the categories.*

<!-- grade_id: cell-827c4d9a4602c0d5, points: 0, grade: false, solution: true -->
```python
pantry = {'pasta': 3, 'garlic': 4,'sauce': 2,
          'basil': 2, 'salt': 3, 'olive oil': 3,
          'rice': 3, 'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 1, 'onions': 1, 'mushrooms': 3,
          'broccoli': 2, 'butter': 2,'pickles': 6, 'milk': 2,
          'chia seeds': 5}

ingredient_categories = {'dairy':['butter', 'milk', 'cheese'], 
                         'produce':['potatoes', 'onions', 'broccoli', 'mushrooms'], 
                         'spices': ['salt', 'pepper', 'paprika']}

category_counts = dict()
for category, ingredients in ingredient_categories.items():
    category_counts[category] = 0
    
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-f4c4f1adf9a87b43, points: 1, grade: true, solution: false -->
```python
t.test_4b(category_counts)
```

## 5. Functions 

**Question 5(a)** <br> {points: 1} 

Remember Arzan from **Question 3(a)** who wanted categories of his heart rate? Well, he is back with another heart rate! 

There is a good chance he may come back again, so to avoid making a new loop every time he returns, a function would be a good way to avoid repetitively writing loops for him. 


Convert the code from **Question 3(b)** into a function, so we don't need to write a lot of code when he comes back. 
Name your function `pulse_categorizer()` and make sure it takes in a numeric value and returns a single heart rate category. 

Use your new function on Arzan's new heart rate stored in the variable `arzan_pulse` and save it in an object named `heartrate_label2`.

<!-- grade_id: cell-ae1be904eb086b18, points: 0, grade: false, solution: true -->
```python
arzan_pulse = 55 

heartrate_label2 = None

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-220ec1ae476bf870, points: 1, grade: true, solution: false -->
```python
t.test_5a(heartrate_label2)
```

**Question 5(b)** <br> {points: 2} 

As we expected, Arzan is back with more even heart rates (we predicted well!). He has brought 5 pulses that he needs categorizing that are stored in a list named `arzan_pulse`.    
Make a loop iterating over `arzan_pulse` that uses your function from **Question 5(a)** to categorize each heart rates and append each heart rate category to a list named `multi_labels`. 

<!-- grade_id: cell-31521646d423b449, points: 0, grade: false, solution: true -->
```python
arzan_pulse = [55, 69, 73, 91, 78]

multi_labels = list()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-e83eb16eca5a5d66, points: 2, grade: true, solution: false -->
```python
# check that the variable exists
assert 'multi_labels' in globals(
), "Please make sure that your solution is named 'multi_labels'"

# This test has been intentionally hidden. It will be up to you to decide if your solution
# is sufficiently good.
```

**Question 5(c)** <br> {points: 1} 

Rupi has come to you asking for the same tool that you created for Hussein in **Question 2(f)**. This was a tool that for every item in a recipe, checked Hussein's pantry and added the ingredients and quantity that Hussein was short on to a shopping list (actually it was a dictionary named `shopping_list`) . 

Repeat **Question 2(f)** for the same `meal_recipe` but this time apply it to Rupi's `fridge` instead of Hussein's `pantry`.  Add the missing ingredients with their quantity to the dictionary `rupis_list`. 

<!-- grade_id: cell-77f4568449330436, points: 0, grade: false, solution: true -->
```python
fridge = {'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 4, 'ketchup': 1, 'mushrooms': 3,
          'vanilla': 1, 'butter': 6,'coconut': 6,
          'cream': 2
          }

meal_recipe = {'pasta': 2, 'garlic': 2, 'sauce': 3,
          'basil': 4, 'salt': 1, 'pepper': 2,
          'olive oil': 2, 'onions': 2, 'mushrooms': 6}


rupis_list = dict()

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-95b259f84f9805a0, points: 1, grade: true, solution: false -->
```python
t.test_5c(rupis_list)
```

**Question 5(d)** <br> {points: 1} 

In the last question, did you notice that you copied/rewrote code, that you've written before? It likely felt a little repetitive. If you find yourself copying or rewriting code, there is a good chance you are in violation of the DRY principle. 


Convert the code from  **Question 5(c)** so that given any pantry or recipe, the function returns the ingredients and their quantities that need to be purchased in a dictionary. 
Here are the function specifications: 

- The function should be named `recipe_shopping_list`. 
- It should have 2 input arguments: 
    - A dictionary containing current items and quantity in their possession and, 
    - A dictionary of recipe ingredients and the amount needed for the recipe. 
- The function should return a dictionary of items and the quantity of which is needed to be purchased.  

<!-- grade_id: cell-c09c338c69145403, points: 0, grade: false, solution: true -->
```python
# Use this as Pantry
dict1 = {'bread': 5, 'peanut butter': 3,
          'flour': 7}

# Use this as Recipe
dict2 = {'eggs': 2, 'vanilla': 2, 'sugar': 3,
                'flour': 2}

# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-15c2554acf286c12, points: 1, grade: true, solution: false -->
```python
t.test_5d(recipe_shopping_list)
```

**Question 5(e)** <br> {points: 1} 

Use the function `recipe_shopping_list` from above to obtain the shopping list of items needed to make the following banana bread recipe given the ingredients in `bakery`.
Save the result in an object named `banana_bread_list`.

<!-- grade_id: cell-16ee00ceb9e97c17, points: 0, grade: false, solution: true -->
```python
bakery = {'bread': 3, 'peanut butter': 1,
          'flour': 1, 'eggs': 4, 'ketchup': 1, 'mushrooms': 3,
          'vanilla': 1, 'butter': 6,'coconut': 6,
          'cream': 2}

banana_bread = {'eggs': 2, 'vanilla': 2, 'sugar': 3,
                'flour': 2, 'butter': 2,
                'cream': 3, 'banana': 2, 'baking soda': 2}

banana_bread_list = None 
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

<!-- grade_id: cell-3c303a06112884aa, points: 1, grade: true, solution: false -->
```python
t.test_5e(banana_bread_list)
```

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.  

## Attributions
- MDS DSCI 511 - Programming for Data Science - [MDS's GitHub website](https://ubc-mds.github.io/course-descriptions/DSCI_511_prog-dsci/) 

## Module Debriefing

If this video is not showing up below, click on the cell and click the  button in the toolbar above.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('RXHwJK90FCI', width=854, height=480)
```