# MCL-DSCI-011 Programming in Python for Data Science 

# Assignment 1: Introduction to Jupyter Notebooks, Python and Pandas

Welcome to MCL-DSCI 011: Programming in Python for Data Science!  

After each module you will complete a module assignment like this one. Before we get started, there are some administrative details.

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are 2 videos that explain:

1. [How to save your Jupyter notebook](https://youtu.be/0aoLgBoAUSA) and,       
2. [How to answer a question in a Jupyter notebook assignment](https://youtu.be/7j0WKhI3W4s).     

If they are not showing up below, click on the cell and click the ▶ button in the toolbar above.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('0aoLgBoAUSA', width=854, height=480)
```

```python
from IPython.display import YouTubeVideo
YouTubeVideo('7j0WKhI3W4s', width=854, height=480)
```

### Assignment Learning Goals:

After completing module 1's assignment you will be able to:

- Describe the components of a Dataframe.
- Read a standard `.csv` file using Pandas [`pd.read_csv()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html).
- Load the `pandas` library into Python.
- Demonstrate indexing and slicing with `.loc[]` and `.iloc[]`.
- Demonstrate Selecting columns of a dataframe using `df[]` notation.
- Obtain values from a dataframe using `.loc[]`.
- Sort a dataframe using [`.sort_values()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.sort_values.html).
- Create simple summary statistics using [`.describe()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.describe.html).
- Construct simple visualizations using [Altair](https://altair-viz.github.io/).
- Create a `.csv` file from a dataframe using [`.to_csv()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.to_csv.html).


In this first worksheet you will also learn how to test the answers you write in this worksheet to assess if you answered questions correctly before your assignment is collected.

This assignment covers parts of [Module 1](https://prog-learn.mds.ubc.ca/en/module1) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` with your completed code and answers then proceed to run the cell!


Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct.

You may notice that we are loading new libraries that you didn't learn during the module. Don't worry, these are needed to properly create and maintain this assignment. 

```python
# Import libraries needed for this lab
import test_assignment1 as t
from hashlib import sha1
import sys
import altair as alt
```

## 1. Jupyter Notebooks
This app is called a Jupyter notebook. A notebook is a place to write programs and view their results.

### 1.1. Text Cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You don't need to learn Markdown, but you might want to.

After you edit a text cell, click the "run cell" button at the top that looks like ▶| to confirm any changes.

**Question 1(a)**   
This paragraph is in its own text cell.  Try editing it so that all of the sentences following this one are deleted, then click the "run cell" ▶| button .  This sentence, for example, should be deleted.  So should this one.

### 1.2. Code Cells
Other cells contain code in the Python language. Running a code cell will execute all of the code it contains.

To run the code in a cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press Run ▶| or hold down the `shift` key and press `return` or `enter`.

Try running the next cell:

```python
"Hello, World!"
```

The above code cell contains a single line of code, but cells can also contain multiple lines of code. When you run a cell, the lines of code are executed in the order in which they appear but only the last line is displayed. 

```python
"First this line is executed but not displayed,"
"and then this line is executed and displayed"
```

**Question 1(b)**    
Change the cell above so that it executes this:

    This line does not get printed,
    same with this line, 
    but this line does!

*Hint:* If you're stuck for more than a few minutes, try talking to a TA during office hours.  That's a good idea for any assignment problem if you get stuck for too long.

### 1.3. Writing Jupyter Notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar. You can also click to the left of a cell so it has a blue outline and hold  `b`. It'll start out as a code cell.  You can change it to a text cell by clicking inside inside the cell so it's highlighted, clicking the drop-down box next to the restart (▶▶) button in the menu bar, and choosing "Markdown". The 'shortcut' keys for this is to make the cell have a blue outline again (by clicking the side of the cell) and press `m` for a markdown/text cell or `y` for a code cell. 

**Question 1(c)**     
Add a code cell below this one.  Write code in it that prints out:
   
    A whole new code cell!

Run your cell to verify that it works.

### 1.4. Errors
Python is a language and, like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them fairly quickly and gain reasonable proficiency with the language in a few months.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes (everyone who writes code does, even the experts!).  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

Below, we disect and explain what is going on in the outputted above. 

![error-message.png](static/error-message.png)

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask for help on Piazza.)

Try to fix the code above so that you can run the cell and see the intended message instead of an error.

### 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
3. If that doesn't help, restart your server. First, save your work by clicking **File** at the top left of your screen, then **Save and Checkpoint**. Next, click **Control Panel** at the top right. Choose **Stop My Server** to shut it down, then **My Server** to start it back up. Then, navigate back to the notebook you were working on.


Like most computer related issues we encourage you to check to make sure your code is running as planned by restarting your kernel often and running your cells again. It's recommended that you take this approach before submitting!


### 1.6. Submitting your work
All assignments in the course will be distributed as notebooks like this one. You will complete your work in this notebook and we will grade it. We use a system called [nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html#) that checks your work.

## 2. Numbers

Quantitative information arises everywhere in data science. In addition to commands to print out lines, our Python code can represent numbers and methods of combining numbers. The expression `3.2500` evaluates to the number 3.25. (Run the cell and see.)

```python
3.2500
```

### 2.1. Arithmetic
The line in the next cell subtracts.  Its value is what you'd expect.  Run it.

```python
2.0 - 1.5
```

Same with the cell below. Run it.

```python
2 * 3
```

Many basic arithmetic operations are built in to Python.  [This webpage](https://en.wikibooks.org/wiki/Python_Programming/Basic_Math) describes all the arithmetic operators used in the course.  You can refer back to this page as you need. 

## 3. Names
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

Similarly, an effective strategy for writing code is to define names for data as we compute it, like a lawyer would define terms for complex ideas at the start of a legal document to simplify the rest of the writing.

In Python, we do this with *objects*. An object has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
answer = 3 * 2 + 4
```

When you run that cell, Python first evaluates the first line.  It computes the value of the expression `3 * 2 + 4`, which is the number 10.  Then it gives that value the name `answer`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `answer`:

```python
answer
```

We can name our objects anything we'd like. Above we called it `answer`, but we could have named it `value`, `data` or anything else we desired. A good rule of thumb is to name it something that has meaning to a human as it relates to what we are trying to accomplish with our  code.

**Question 3(a)**   
Enter a new code cell under this cell and create another object using `= 3 * 2 + 4` with the name `practice_arithmetic`. 

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355 / 113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
bimonthly_salary = 840
monthly_salary = 2 * bimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).   Names can't contain spaces since spaces are used to separate pieces of code from each other, nor can they contain other special characters such as `+`, `#`, `$`, `%`, `^` since some characters have special roles in python. Take `#` for example, this character specifies a comment within written code. 

Other than those rules, what you name something doesn't matter to Python.  For example, the next cell does the same thing as the above cell, except everything has a different name:


Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers. The first character can't be a number (otherwise a name might look like a number). And names can't contain spaces, since spaces are used to separate pieces of code from each other." -> they can't contain other characters like + either

```python
a = 840
b = 2 * a
c = 12
d = c * b
d
```

**However**, names are very important for making your code *readable* to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does. 

There is also cultural style associated with different programming languages. In the modern Python style, object names should use only lowercase letters, numbers, and `_`. Underscores (`_`) are typically used to separate words within a name (*e.g.*, `answer_one`).

### 3.1. Comments
Below you see lines like this in code cells:

    # Test cell; please do not change!

That is called a *comment*.  It doesn't make anything happen in Python; Python ignores anything on a line after a #.  Instead, it's there to communicate something about the code to you, the human reader.  Comments are extremely useful and can help increase how readable our code is.

![](http://imgs.xkcd.com/comics/future_self.png)

**Question 3(b)** <br> {points: 1}

Assign the name `seconds_in_an_hour` to the number of seconds in an hour. You should do this in two steps. In the first, you calculate the number of seconds in a minute and assign that number the name `seconds_in_a_minute`. Next you should calculate the number of seconds in an hour and assign that number the name `seconds_in_an_hour.`  

<!-- grade_id: cell-cc6e31413919a5ae, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
seconds_in_a_minute = 60
seconds_in_an_hour = seconds_in_a_minute * 60
### END SOLUTION

# We've put this line in this cell so that it will print
# the value you've given to seconds_in_an_hour when you
# run it.  You don't need to change this.
seconds_in_an_hour
```

### 3.2. Checking your code

Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. We imported the code below earlier at the begining of the assignment and you will see a variation of this cell at the very top of all future assignments.

```python
import test_assignment1 as t
```

Below is an example of a test cell for Question 3(b) above (assesses whether you have assigned `seconds_in_an_hour` correctly). If you haven't, this test will tell you that your _"Solution is incorrect"_. Resist the urge to just copy it, and instead try to adjust your expression. (Sometimes the tests will give hints about what went wrong...)

<!-- grade_id: cell-0a5f8cd77d5b9fc0, points: 1, grade: true, solution: false -->
```python
t.test_3b(seconds_in_a_minute,seconds_in_an_hour)
```

For this first question we'll provide you the solution:

```python
# Calculate the number of seconds in an hour.

#SOLUTION:
seconds_in_a_minute = 60
seconds_in_an_hour = seconds_in_a_minute * 60

# We've put this line in this cell so that it will print
# the value you've given to seconds_in_an_hour when you
# run it.  You don't need to change this.
seconds_in_an_hour
```

```python
t.test_3b(seconds_in_a_minute,seconds_in_an_hour)
```

## 4. Calling Functions/Methods and Attributes

The most common way to combine or manipulate values in Python is by calling functions. Python comes with many built-in functions and attributes that perform common operations.

In the module, we explored examples of functions and methods such as `describe()` and `sort_values()`.
Here, we'll demonstrate using another method `upper()` that converts text to uppercase:

```python
greeting = "Why, hello there!"
greeting
```

```python
greeting.upper()
```

**Question 4(a)** <br> {points: 1} 

Use the method `lower()` to change all the characters in the following movie title to lower case text:
"The House with a Clock in Its WaLls" 
Assign the lower case text the name `title`.

<!-- grade_id: cell-2339ec0ab12fd9f6, points: 0, grade: false, solution: true -->
```python
title = "The House with a Clock in Its WaLls"
### BEGIN SOLUTION
title.lower()
### END SOLUTION

```

<!-- grade_id: cell-213a9370c52a3857, points: 1, grade: true, solution: false -->
```python
assert 'title' in locals(), "Your answer should be named 'title'"

t.test_4a(title)
```

### 4.1. Multiple Arguments
Some functions take multiple arguments, separated by commas. For example, the built-in `max` function returns the maximum argument passed to it.

```python
max(2, 15, 4, 7)
```

**Question 4(b)** <br> {points: 1}

Use the `min` function to find the minumum value of the numbers in the cell above.

Assign the value to an object called `smallest`.

<!-- grade_id: cell-9b48c973a6f73f6d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
smallest = min(2, 15, 4, 7)
smallest
### END SOLUTION
```

<!-- grade_id: cell-491b8792be2952b8, points: 1, grade: true, solution: false -->
```python
assert 'smallest' in locals(), "Your answer should be named 'smallest'"

t.test_4b(smallest)
```

## 5. Importing Packages - Pandas 

Python has many built-in functions, but we can also use functions that are stored within packages created by other Python users. We discussed in the module about needing `pandas` to work with data frames.


**Question 5(a)** <br> {points: 1}

Let's import `pandas` as as `pd`. 

<!-- grade_id: cell-6a1897ad6c6de7ee, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import pandas as pd
### END SOLUTION
```

<!-- grade_id: cell-688755597ea2135c, points: 1, grade: true, solution: false -->
```python
t.test_5a(dir())
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-3-191167b88bc4> in <module>
----> 1 t.test_5a(dir())

TypeError: test_5a() takes 0 positional arguments but 1 was given
```

## 6. Looking for Help

No one, even experienced, professional programmers remember what every function does, nor do they remember every possible function argument/option. So both experienced and new programmers (like you!) need to look things up A LOT! 

### Help Files

#### Functions

One of the most efficient places to look for help on how a function works is the Python help files. 


Let’s say we wanted to pull up the help file for the `max()` function. We can do this by typing a question mark in front of the function we want to know more about.
Run the cells below to find out more about `max` and `pd.read_csv`.


```python
?max
```

```python
?pd.read_csv
```

At the very top of the file, you will see the function itself and all of the possible arguments it takes (we will go over this in Assignment 3). Next is a description of what the function does. You’ll find that the most helpful sections on this page are “Usage”, “Arguments” and "Examples". 

- **Docstring:** gives you an idea of how you would use the function when coding--what the syntax would be and how the function itself is structured. 
- **Parameters** the possible inputs to the function. Often the  “Parameters” sections don’t provide you with step by step instructions, because there are so many different ways that a person can incorporate a function into their code. Instead, they provide users with a general understanding as to what the function could do and parts that could be added. At the end of the day, the user must interpret the help file and figure out how best to use the functions and which parts are most important to include for their particular task. 
- The **Returns** explains what to expect as an output. 
- The **Examples** section is often the most useful part of the help file as it shows how a function could be used with real data. It provides a skeleton code that the users can work off of.
- Sometimes there is a **See Also** section which may suggest similar functions that could be of use to the user. 

**Question 6(a)** Multiple Choice:
<br> {points: 1}

Use `?pd.read_csv()` to answer the multiple choice question below. To answer the question, assign the letter associated with the correct answer to a variable in the the code cell below:

Which statement below is accurate?

A. `pd.read_csv` is useful for reading comma-separated values (csv) file into DataFrame. 

B. It can accept a possible parameter of warnings=True

C. The parameter `delimiter` is an Alias for the parameter `squeeze`

D. `pd.read_csv` is perfect for reading a table of fixed-width formatted lines into DataFrame.

*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer6_a`.*

<!-- grade_id: cell-1d73b03a76f49bd1, points: 0, grade: false, solution: true -->
```python
answer6_a = None

### BEGIN SOLUTION
answer6_a = "A"
### END SOLUTION

```

<!-- grade_id: cell-8b8b2076f270144a, points: 2, grade: true, solution: false -->
```python
assert 'answer6_a' in locals(), "Your answer should be named 'answer6_a'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS

def hidden_1(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    answer = answer.upper()
    assert sha1(str(answer).encode('utf8')).hexdigest(
    ) == '6dcd4ce23d88e2ee9568ba546c007c63d9131c1b', "Your answer is incorrect." # we hid the answer to the test here so you can't see it, but we can still run the test
    return ("Success")
hidden_1(answer6_a)
### END HIDDEN TESTS
```

#### Methods and Attributes

We will explain in more detail the difference between functions, methods and attributes later in this course, however until then you will need to simply know how to use them and potentially find documentation on them in order to do so.
If you attempt to use `?` on a method or an attribute, it will return `Object not found.`. This means that in order to obtain any help files for the verb, you will need to include an object with it. 
We will use the canucks dataset for our example.

```python
df = pd.read_csv('data/canucks.csv')
```

If we attempted to get more information on the verb `head` as we did before, we would run into an error. 

```python
?head
```

To obtain the documentation for `head`, we need to include an object, which in this case, is the dataframe `df`.

```python
?df.head
```

**Question 6(b)** Multiple Choice:
<br> {points: 1}

Use `?` to answer the multiple choice question below. To answer the question, assign the letter associated with the correct answer to a variable in the the code cell below:

Which statement below is accurate regarding the method `.mean` ?

A. The default for the `skipna` argument is `False`

B. It only returns Dataframes

C. The mean of the values for the requested axis is returned

D. There is an argument named `columns` with a default setting of `None` 

Answer in the cell below using the uppercase letter associated with your answer. Place your answer between "", assign the correct answer to an object called `answer6_b`.

<!-- grade_id: cell-324e41ba9b74f760, points: 0, grade: false, solution: true -->
```python
answer6_b = None

### BEGIN SOLUTION
answer6_b = "C"
### END SOLUTION
```

<!-- grade_id: cell-5bcee13ca6500cb7, points: 1, grade: true, solution: false -->
```python
t.test_6b(answer6_b)
```

Beyond the Docstring help files, there are many resources that you can use to find help. [Stack Overflow](https://stackoverflow.com/), an online forum, is a great place to go and ask questions such as how to perform a complicated task in a number of different languages including Python or why a specific error message is popping up. Oftentimes, a previous user will have already asked your question of interest and received helpful advice from fellow users.

## 7. Gentle Pandas 

Now that we have learned a little about Jupyter notebooks and Python, let's load a real dataset and explore it. As we do this we will learn more about key data loading, wrangling and visualization functions in Python. 


### 7.0. Data about Runners!

Researchers Vickers and Vertosick performed [a study in 2016](https://bmcsportsscimedrehabil.biomedcentral.com/articles/10.1186/s13102-016-0052-y) that aimed to identify what factors had a relationship with race performance of recreational runners so that they could better predict future 5 km, 10 km and marathon race times for individual runners. Such predictions (and knowing what drives these predictions) can help runners by suggesting changes they could make to modifiable factors, such as training, to help them improve race time. Unmodifiable factors that contribute to the prediction, such as age or sex, allow for fair comparisons to be made between different runners.


To make a prediction, Vickers and Vertosick performed a large survey and were able to obtain 2,497 responses. The survey included questions that allowed them to collect a data set that included: 
- age,
- sex,
- body mass index (BMI),
- whether they are an edurance runner or speed demon,
- what type of shoes they wear,
- what type of training they do,
- race time for 2-3 races they completed in the last 6 months,
- self-rated fitness for each race,
- and race difficulty for each race.

***Note: many popular datasets have sex as a feature where the possible values are male and female. This representation reflects how the data were collected and is not meant to imply that, for example, gender is binary.***

Let's now use this data to explore a question we might be interested in - is there a relationship between 10 km race time and body mass index (BMI) for men runners in this data set. This is an exploratory data analysis question because we stated we are looking for a relationship between measurements within the single data set we have and are not interested in yet interpreting beyond it. We can answer this question by visualizing the data as a scatter plot.

Some useful things to know about this dataset: 
- The runners are sorted with males at the top of the dataframe. 
- There are 1141 males racing and 692 females. 

Let's go ahead and start our exploration to answer their question. 

**Question 7(a)** <br> {points: 1}   
Load in the data named `marathon_small.csv` and save it in an object named `all_runners`.

<!-- grade_id: cell-1151a4ae5e2c24a0, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
all_runners = pd.read_csv("marathon_small_sorted.csv")
all_runners
### END SOLUTION
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-486555bf5e13> in <module>
      1 ### BEGIN SOLUTION
----> 2 all_runners = pd.read_csv("marathon_small_sorted.csv")
      3 all_runners
      4 ### END SOLUTION

NameError: name 'pd' is not defined
```

<!-- grade_id: cell-ebf793ef459f7ff3, points: 1, grade: true, solution: false -->
```python
t.test_7a(all_runners)
```

**Question 7(b)** <br> {points: 1}     
Subset the data we are interested in. In this case we want all the male runners and all the columns. Save this dataframe `male_runners`.    
_This dataframe is particularly convenient because the first 1141 rows are male_

<!-- grade_id: cell-2865faf54e339127, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
male_runners = all_runners.iloc[0:1141, :]
male_runners
### END SOLUTION

```

<!-- grade_id: cell-6d8376097f1c7f37, points: 1, grade: true, solution: false -->
```python
t.test_7b(male_runners)
```

**Question 7(c)** <br> {points: 1}     
Sort the `male_runners` dataframe from fastest to slowest 10km time in minutes and label the dataframe as `sorted_runners`. 

<!-- grade_id: cell-5128f317b0c83700, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
sorted_runners = male_runners.sort_values("km10_time_minutes")
sorted_runners
### END SOLUTION

```

<!-- grade_id: cell-1c1ed6887638c576, points: 2, grade: true, solution: false -->
```python
assert 'sorted_runners' in locals(), "Your answer should be named 'sorted_runners'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS
def hidden_2(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    assert isinstance(answer, type(pd.DataFrame())), str(answer) + \
        " should be a pandas dataframe"
    assert len(answer) == 1141, "the number of rows in the dataframe should be 1141"
    assert len(answer.columns) == 6, "the number of columns in the dataframe should be 6. \n Did you load it properly?"
    assert sorted(list(answer.columns.values)) == ['age',
                                                   'bmi',
                                                   'km10_time_minutes',
                                                   'km10_time_seconds',
                                                   'km5_time_seconds',
                                                   'sex'], "At least one column name is not matching the original dataframe"
    assert answer["km10_time_minutes"][0:5].sum() == 147.8,  "Your dataframe has the wrong summed values for the fastest 5 runners."
    return ("Success")

hidden_2(sorted_runners)
### END HIDDEN TESTS
```

**Question 7(d)** <br> {points: 1}   
Find some summary statistics among male runners using `pd.describe()` and obtaining the numerical columns only. Save it as `male_stats` and display it. 

<!-- grade_id: cell-2523b4de0466122f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
male_stats = male_runners.describe()
male_stats
### END SOLUTION

```

<!-- grade_id: cell-fd67c2720c342627, points: 1, grade: true, solution: false -->
```python
t.test_7d(male_stats)
```

**Question 7(e)** 

Use `.loc[]` to obtain the values from the `male_stats` dataframe to answer the following questions. 

**Question 7(e)(i)** <br> {points: 1}   
Using `.loc[]`, obtain the value in the `male_stats`
What is the mean age of the male runners ? Save it as object `ans_7_e_i`

<!-- grade_id: cell-c354bf548da124ad, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ans_7_e_i = male_stats.loc['mean', 'age']
ans_7_e_i
### END SOLUTION
```

<!-- grade_id: cell-31966384da74b315, points: 1, grade: true, solution: false -->
```python
t.test_7e_i(ans_7_e_i)
```

**Question 7(e)(ii)** <br> {points: 1}   
What is the largest BMI of the male runners? Save it as object `ans_7_e_ii`


<!-- grade_id: cell-95254f3cf942d659, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ans_7_e_ii = male_stats.loc['max', 'bmi']
ans_7_e_ii
### END SOLUTION
```

<!-- grade_id: cell-a973967f50619fc0, points: 2, grade: true, solution: false -->
```python
assert 'ans_7_e_ii' in locals(), "Your answer should be named 'ans_7_e_ii'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS

def hidden_3(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    assert isinstance(answer, float), str(answer) + \
        " Are you putting quotes around a number?"
    assert sha1(str(round(answer + 7,2)).encode('utf8')).hexdigest(
    ) == 'e544ffc8c8dc94fcbab67753a5fc03c9360d5478', "Your answer is incorrect. You should indicate the index corresponding to the wrongly labled sample"
    return("Success")

hidden_3(ans_7_e_ii)
### END HIDDEN TESTS
```

**Question 7(e)(iii)** <br> {points: 1}   
What is the fastest	5 km time in seconds of the male runners? Save it as object `ans_7_e_iii`


<!-- grade_id: cell-b8381f581cb76166, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ans_7_e_iii = male_stats.loc['min', 'km5_time_seconds']
ans_7_e_iii
### END SOLUTION
```

<!-- grade_id: cell-24eb0f66fa29fe63, points: 1, grade: true, solution: false -->
```python
t.test_7e_iii(ans_7_e_iii)
```

**Question 7(e){iv)** <br> {points: 1}   
a) What is the 50th percentile (also known as the median) age? Save it as object `ans_7_e_iv`

<!-- grade_id: cell-4f9f7ee9dba0b086, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ans_7_e_iv = male_stats.loc['50%', 'age']
ans_7_e_iv
### END SOLUTION
```

<!-- grade_id: cell-6b288a9a2816b432, points: 1, grade: true, solution: false -->
```python
t.test_7e_iv(ans_7_e_iv)
```

**Question 7(f)** <br> {points: 5}   
What is the frequency of each age among male runners? 
- This is a multi-step question. You will need to have multiple objects before obtaining your end object. 
- Use `.value_counts()` and save your table as `male_age_freq`

<!-- grade_id: cell-3d351db39a27720c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
male_ages = male_runners['age']
male_age_freq = male_ages.value_counts()
male_age_freq
### END SOLUTION
```

<!-- grade_id: cell-dac777f80f7185dd, points: 5, grade: true, solution: false -->
```python
t.test_7f(male_age_freq)
```

### Histograms

In the lecture material we practiced making bar charts. But let's try something a little different. Bar charts are used for categorical variables but what if you wanted to see the distribution of the male runners age which is a continuous variable? In this case you would have to plot a histogram. The lucky thing is, it's simpler than bar charts.   
Let's plot the histogram of male runner's age together.

```python
# to plot a histogram, we use the .mark_bar() function
# setting y='count()' tells altair that we would like to count the number of individual in 
# each age group
alt.Chart(male_runners).mark_bar().encode(
    x='age',
    y='count()'
)
```

We can use the `alt.X()` and `alt.Y()` arguments in the `encode` function to provide more control over how we would the plot to look. For example, if we wanted to change the number of bins, we can use the `bin=alt.Bin(maxbins=..)` in `alt.X()` to set a value for the number of bins. Let's change the number of bins from the default to `50` by `bin=alt.Bin(maxbins=50)`. 

```python
# set the number of bins to 50

alt.Chart(male_runners).mark_bar().encode(
    x=alt.X('age', bin=alt.Bin(maxbins=50)),
    y='count()'
)
```

It's always good to provide a title to your plot. It makes it easier for others to have an understanding of what your are trying to convey. We can add a title to a plot using the `.properties(title="")` function. Lets try it out below.

```python
# provide a title

alt.Chart(male_runners).mark_bar().encode(
    x=alt.X('age', bin=alt.Bin(maxbins=50)),
    y=alt.Y('count()')
).properties(title="Here is a title for this plot")
```

Alright, it's your turn to try it now. 

**Question 7(g) Something New !** <br> {points: 5}  

Plot the distribution (histogram) of male runners' BMIs and save it as `hist_bmi`. Use the argument `bin=alt.Bin(maxbins=25)`, change the colour to `lightsteelblue` and add a title called "Distribution of Male Runners BMI".

<!-- grade_id: cell-a8491236ce94bd95, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
hist_bmi = alt.Chart(male_runners).mark_bar(color="lightsteelblue").encode(
    x=alt.X('bmi', bin=alt.Bin(maxbins=25)),
    y=alt.Y('count()')
).properties(title="Distribution of Male Runners BMI")
### END SOLUTION
```

<!-- grade_id: cell-f4fb28121f04ae60, points: 5, grade: true, solution: false -->
```python
t.test_7g(hist_bmi)
```

**Question 7(h) Answering the Question** <br > {points: 5}

Let's try and answer the original question that was asked:   

***Is there a relationship between 10 km race time and body mass index (BMI) for men runners in this data set.***

Plot a scatterplot of our sliced dataframe to see if there is a relationship between 10 km race time and body mass index(BMI) for men runners.

- Use the `mark_circle()` function to create a scatter plot.
- Name your plot `scatter_bmi_10k`
- Change the colour to `slateblue`
- Make the size of the points 12
- Opacity should be 30 % of regular intensity.
- Add a title "Relationship between BMI and 10k running time (mins)"

<!-- grade_id: cell-c3a6760fbd6e2b27, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

scatter_bmi_10k = alt.Chart(male_runners).mark_circle(color="slateblue",size=12,opacity=0.3).encode(
    x=alt.X('bmi'),
    y=alt.Y('km10_time_minutes')
).properties(title="Relationship between BMI and 10k running time (mins)")

### END SOLUTION
```

<!-- grade_id: cell-b794f114d6860fe3, points: 5, grade: true, solution: false -->
```python
t.test_7h(scatter_bmi_10k)
```

**Question 7(i) Answering the Question** <br> {points: 1}  
Does there appear to be a relationship present between BMI and 10 km race time among male runners?
A **positive** relationship is present if when we increase one variable, the other variable tend to increase also. A **negative** relationship is present if increasing one variable has the opposite effect on the other and decreases the other variable.
To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:

A. No, there doesn't appear to be a relationship     

B. Yes, there appears to be a positive relationship      

C. Yes, there appears to be a negative relationship      


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer7_i`.*

<!-- grade_id: cell-61cc577e15771efc, points: 0, grade: false, solution: true -->
```python
answer7_i = None

### BEGIN SOLUTION
answer7_i = "B"
### END SOLUTION
```

<!-- grade_id: cell-3660be164cbdbe7b, points: 1, grade: true, solution: false -->
```python
t.test_7i(answer7_i)
```

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.   

## Attributions
- UC Berkley [Data 8 Public Materials](https://github.com/data-8/data8assets)
- UBC's [DSCI 100 Public Material](https://github.com/UBC-DSCI/dsci-100-assets)

## Module Debriefing

If this video is not showing up below, click on the cell and click the ▶ button in the toolbar above.

```python
from IPython.display import YouTubeVideo
YouTubeVideo('es-iDoCIRP0', width=854, height=480)
```